Human Brain-Wide Activation of Sleep Rhythms

Author response:

The following is the authors’ response to the original reviews

Public Reviews:

Reviewer #1 (Public review):

Wang et al., recorded concurrent EEG-fMRI in 107 participants during nocturnal NREM sleep to investigate brain activity and connectivity related to slow oscillations (SO), sleep spindles, and in particular their co-occurrence. The authors found SO-spindle coupling to be correlated with increased thalamic and hippocampal activity, and with increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the neocortex, especially the medial prefrontal cortex (mPFC). They concluded the brain-wide activation pattern to resemble episodic memory processing, but to be dissociated from task-related processing and suggest that the thalamus plays a crucial role in coordinating the hippocampal-cortical dialogue during sleep.

The paper offers an impressively large and highly valuable dataset that provides the opportunity for gaining important new insights into the network substrate involved in SOs, spindles, and their coupling. However, the paper does unfortunately not exploit the full potential of this dataset with the analyses currently provided, and the interpretation of the results is often not backed up by the results presented. I have the following specific comments.

Thank you for your thoughtful and constructive feedback. We greatly appreciate your recognition of the strengths of our dataset and findings Below, we address your specific comments and provide responses to each point you raised to ensure our methods and results are as transparent and comprehensible as possible. We hope these revisions address your comments and further strengthen our manuscript. Thank you again for the constructive feedback.

(1) The introduction is lacking sufficient review of the already existing literature on EEG-fMRI during sleep and the BOLD-correlates of slow oscillations and spindles in particular (Laufs et al., 2007; Schabus et al., 2007; Horovitz et al., 2008; Laufs, 2008; Czisch et al., 2009; Picchioni et al., 2010; Spoormaker et al., 2010; Caporro et al., 2011; Bergmann et al., 2012; Hale et al., 2016; Fogel et al., 2017; Moehlman et al., 2018; Ilhan-Bayrakci et al., 2022). The few studies mentioned are not discussed in terms of the methods used or insights gained.

We acknowledge the need for a more comprehensive review of prior EEG-fMRI studies investigating BOLD correlates of slow oscillations and spindles. However, these articles are not all related to sleep SO or spindle. Articles (Hale et al., 2016; Horovitz et al., 2008; Laufs, 2008; Laufs, Walker, & Lund, 2007; Spoormaker et al., 2010) mainly focus on methodology for EEG-fMRI, sleep stages, or brain networks, which are not the focus of our study. Thank you again for your attention to the comprehensiveness of our literature review, and we will expand the introduction to include a more detailed discussion of the existing literature, ensuring that the contributions of previous EEG-fMRI sleep studies are adequately acknowledged.

Introduction, Page 4 Lines 62-76

“Investigating these sleep-related neural processes in humans is challenging because it requires tracking transient sleep rhythms while simultaneously assessing their widespread brain activation. Recent advances in simultaneous EEG-fMRI techniques provide a unique opportunity to explore these processes. EEG allows for precise event-based detection of neural signal, while fMRI provides insight into the broader spatial patterns of brain activation and functional connectivity (Horovitz et al., 2008; Huang et al., 2024; Laufs, 2008; Laufs, Walker, & Lund, 2007; Schabus et al., 2007; Spoormaker et al., 2010). Previous EEG-fMRI studies on sleep have focused on classifying sleep stages or examining the neural correlates of specific waves (Bergmann et al., 2012; Caporro et al., 2012; Czisch et al., 2009; Fogel et al., 2017; Hale et al., 2016; Ilhan-Bayrakcı et al., 2022; Moehlman et al., 2019; Picchioni et al., 2011). These studies have generally reported that slow oscillations are associated with widespread cortical and subcortical BOLD changes, whereas spindles elicit activation in the thalamus, as well as in several cortical and paralimbic regions. Although these findings provide valuable insights into the BOLD correlates of sleep rhythms, they often do not employ sophisticated temporal modeling (Huang et al., 2024), to capture the dynamic interactions between different oscillatory events, e.g., the coupling between SOs and spindles.”

(2) The paper falls short in discussing the specific insights gained into the neurobiological substrate of the investigated slow oscillations, spindles, and their interactions. The validity of the inverse inference approach ("Open ended cognitive state decoding"), assuming certain cognitive functions to be related to these oscillations because of the brain regions/networks activated in temporal association with these events, is debatable at best. It is also unclear why eventually only episodic memory processing-like brain-wide activation is discussed further, despite the activity of 16 of 50 feature terms from the NeuroSynth v3 dataset were significant (episodic memory, declarative memory, working memory, task representation, language, learning, faces, visuospatial processing, category recognition, cognitive control, reading, cued attention, inhibition, and action).

Thank you for pointing this out, particularly regarding the use of inverse inference approaches such as “open-ended cognitive state decoding.” Given the concerns about the indirectness of this approach, we decided to remove its related content and results from Figure 3 in the main text and include it in Supplementary Figure 7. We will refocus the main text on direct neurobiological insights gained from our EEG-fMRI analyses, particularly emphasizing the hippocampal-thalamocortical network dynamics underlying SO-spindle coupling, and we will acknowledge the exploratory nature of these findings and highlight their limitations.

Discussion, Page 17-18 Lines 323-332

“To explore functional relevance, we employed an open-ended cognitive state decoding approach using meta-analytic data (NeuroSynth: Yarkoni et al. (2011)). Although this method usefully generates hypotheses about potential cognitive processes, particularly in the absence of a pre- and post-sleep memory task, it is inherently indirect. Many cognitive terms showed significant associations (16 of 50), such as “episodic memory,” “declarative memory,” and “working memory.” We focused on episodic/declarative memory given the known link with hippocampal reactivation (Diekelmann & Born, 2010; Staresina et al., 2015; Staresina et al., 2023). Nonetheless, these inferences regarding memory reactivation should be interpreted cautiously without direct behavioral measures. Future research incorporating explicit tasks before and after sleep would more rigorously validate these potential functional claims.”

(3) Hippocampal activation during SO-spindles is stated as a main hypothesis of the paper - for good reasons - however, other regions (e.g., several cortical as well as thalamic) would be equally expected given the known origin of both oscillations and the existing sleep-EEG-fMRI literature. However, this focus on the hippocampus contrasts with the focus on investigating the key role of the thalamus instead in the Results section.

We appreciate your insight regarding the relative emphasis on hippocampal and thalamic activation in our study. We recognize that the manuscript may currently present an inconsistency between our initial hypothesis and the main focus of the results. To address this concern, we will ensure that our Introduction and Discussion section explicitly discusses both regions, highlighting the complementary roles of the hippocampus (memory processing and reactivation) and the thalamus (spindle generation and cortico-hippocampal coordination) in SO-spindle dynamics.

Introduction, Page 5 Lines 87-103

“To address this gap, our study investigates brain-wide activation and functional connectivity patterns associated with SO-spindle coupling, and employs a cognitive state decoding approach (Margulies et al., 2016; Yarkoni et al., 2011)—albeit indirectly—to infer potential cognitive functions. In the current study, we used simultaneous EEG-fMRI recordings during nocturnal naps (detailed sleep staging results are provided in the Methods and Table S1) in 107 participants. Although directly detecting hippocampal ripples using scalp EEG or fMRI is challenging, we expected that hippocampal activation in fMRI would coincide with SO-spindle coupling detected by EEG, given that SOs, spindles, and ripples frequently co-occur during NREM sleep. We also anticipated a critical role of the thalamus, particularly thalamic spindles, in coordinating hippocampal-cortical communication.

We found significant coupling between SOs and spindles during NREM sleep (N2/3), with spindle peaks occurring slightly before the SO peak. This coupling was associated with increased activation in both the thalamus and hippocampus, with functional connectivity patterns suggesting thalamic coordination of hippocampal-cortical communication. These findings highlight the key role of the thalamus in coordinating hippocampal-cortical interactions during human sleep and provide new insights into the neural mechanisms underlying sleep-dependent brain communication. A deeper understanding of these mechanisms may contribute to future neuromodulation approaches aimed at enhancing sleep-dependent cognitive function and treating sleep-related disorders.”

Discussion, Page 16-17 Lines 292-307

“When modeling the timing of these sleep rhythms in the fMRI, we observed hippocampal activation selectively during SO-spindle events. This suggests the possibility of triple coupling (SOs–spindles–ripples), even though our scalp EEG was not sufficiently sensitive to detect hippocampal ripples—key markers of memory replay (Buzsáki, 2015). Recent iEEG evidence indicates that ripples often co-occur with both spindles (Ngo, Fell, & Staresina, 2020) and SOs (Staresina et al., 2015; Staresina et al., 2023). Therefore, the hippocampal involvement during SO-spindle events in our study may reflect memory replay from the hippocampus, propagated via thalamic spindles to distributed cortical regions.

The thalamus, known to generate spindles (Halassa et al., 2011), plays a key role in producing and coordinating sleep rhythms (Coulon, Budde, & Pape, 2012; Crunelli et al., 2018), while the hippocampus is found essential for memory consolidation (Buzsáki, 2015; Diba & Buzsá ki, 2007; Singh, Norman, & Schapiro, 2022). The increased hippocampal and thalamic activity, along with strengthened connectivity between these regions and the mPFC during SO-spindle events, underscores a hippocampal-thalamic-neocortical information flow. This aligns with recent findings suggesting the thalamus orchestrates neocortical oscillations during sleep (Schreiner et al., 2022). The thalamus and hippocampus thus appear central to memory consolidation during sleep, guiding information transfer to the neocortex, e.g., mPFC.”

(4) The study included an impressive number of 107 subjects. It is surprising though that only 31 subjects had to be excluded under these difficult recording conditions, especially since no adaptation night was performed. Since only subjects were excluded who slept less than 10 min (or had excessive head movements) there are likely several datasets included with comparably short durations and only a small number of SOs and spindles and even less combined SO-spindle events. A comprehensive table should be provided (supplement) including for each subject (included and excluded) the duration of included NREM sleep, number of SOs, spindles, and SO+spindle events. Also, some descriptive statistics (mean/SD/range) would be helpful.

We appreciate your recognition of our sample size and the challenges associated with simultaneous EEG-fMRI sleep recordings. We acknowledge the importance of transparently reporting individual subject data, particularly regarding sleep duration and the number of detected SOs, spindles, and SO-spindle events. To address this, we will provide comprehensive tables in the supplementary materials, contains descriptive information about sleep-related characteristics (Table S1), as well as detailed information about sleep waves at each sleep stage for all 107 subjects(Table S2-S4), listing for each subject:(1)Different sleep stage duration; (2)Number of detected SOs; (3)Number of detected spindles; (4)Number of detected SO-spindle coupling events; (5)Density of detected SOs; (6)Density of detected spindles; (7)Density of detected SO-spindle coupling events.

However, most of the excluded participants were unable to fall asleep or had too short a sleep duration, so they basically had no NREM sleep period, so it was impossible to count the NREM sleep duration, SO, spindle, and coupling numbers.

Supplementary Materials, Page 42-54, Table S1-S4

(5) Was the 20-channel head coil dedicated for EEG-fMRI measurements? How were the electrode cables guided through/out of the head coil? Usually, the 64-channel head coil is used for EEG-fMRI measurements in a Siemens PRISMA 3T scanner, which has a cable duct at the back that allows to guide the cables straight out of the head coil (to minimize MR-related artifacts). The choice for the 20-channel head coil should be motivated. Photos of the recording setup would also be helpful.

Thank you for your comment regarding our choice of the 20-channel head coil for EEG-fMRI measurements. We acknowledge that the 64-channel head coil is commonly used in Siemens PRISMA 3T scanners; however, the 20-channel coil was selected due to specific practical and technical considerations in our study. In particular, the 20-channel head coil was compatible with our EEG system and ensured sufficient signal-to-noise ratio (SNR) for both EEG and fMRI acquisition. The EEG electrode cables were guided through the lateral and posterior openings of the head coil, secured with foam padding to reduce motion and minimize MR-related artifacts. Moreover, given the extended nature of nocturnal sleep recordings, the 20-channel coil allowed us to maintain participant comfort while still achieving high-quality simultaneous EEG-fMRI data.

We have made this clearer in the revised manuscript.

Methods, Page 20 Lines 385-392

“All MRI data were acquired using a 20-channel head coil on a research-dedicated 3-Tesla Siemens Magnetom Prisma MRI scanner. Earplugs and cushions were provided for noise protection and head motion restriction. We chose the 20-channel head coil because it was compatible with our EEG system and ensured sufficient signal-to-noise ratio (SNR) for both EEG and fMRI acquisition. The EEG electrode cables were guided through the lateral and posterior openings of the head coil, secured with foam padding to reduce motion and minimize MR-related artifacts. Moreover, given the extended nature of nocturnal sleep recordings, the 20-channel coil helped maintain participant comfort while still achieving high-quality simultaneous EEG-fMRI data.”

(6) Was the EEG sampling synchronized to the MR scanner (gradient system) clock (the 10 MHz signal; not referring to the volume TTL triggers here)? This is a requirement for stable gradient artifact shape over time and thus accurate gradient noise removal.

Thank you for raising this important point. We confirm that the EEG sampling was synchronized to the MR scanner’s 10 MHz gradient system clock, ensuring a stable gradient artifact shape over time and enabling accurate artifact removal. This synchronization was achieved using the standard clock synchronization interface of the EEG amplifier, minimizing timing jitter and drift. As a result, the gradient artifact waveform remained stable across volumes, allowing for more effective artifact correction during preprocessing. We appreciate your attention to this critical aspect of EEG-fMRI data acquisition.

We have made this clearer in the revised manuscript.

Methods, Page 19-20 Lines 371-383

“EEG was recorded simultaneously with fMRI data using an MR-compatible EEG amplifier system (BrainAmps MR-Plus, Brain Products, Germany), along with a specialized electrode cap. The recording was done using 64 channels in the international 10/20 system, with the reference channel positioned at FCz. In order to adhere to polysomnography (PSG) recording standards, six electrodes were removed from the EEG cap: one for electrocardiogram (ECG) recording, two for electrooculogram (EOG) recording, and three for electromyogram (EMG) recording. EEG data was recorded at a sample rate of 5000 Hz, the resistance of the reference and ground channels was kept below 10 kΩ, and the resistance of the other channels was kept below 20 kΩ. To synchronize the EEG and fMRI recordings, the BrainVision recording software (BrainProducts, Germany) was utilized to capture triggers from the MRI scanner. The EEG sampling was synchronized to the MR scanner’s 10 MHz gradient system clock, ensuring a stable gradient artifact shape over time and enabling accurate artifact removal. This was achieved via the standard clock synchronization interface of the EEG amplifier, minimizing timing jitter and drift.”

(7) The TR is quite long and the voxel size is quite large in comparison to state-of-the-art EPI sequences. What was the rationale behind choosing a sequence with relatively low temporal and spatial resolution?

We acknowledge that our chosen TR and voxel size are relatively long and large compared to state-of-the-art EPI sequences. This decision was made to optimize the signal-to-noise ratio (SNR) and reduce susceptibility-related distortions, which are particularly critical in EEG-fMRI sleep studies where head motion and physiological noise can be substantial. A longer TR allowed us to sample whole-brain activity with sufficient coverage, while a larger voxel size helped enhance BOLD sensitivity and minimize partial volume effects in deep brain structures such as the thalamus and hippocampus, which are key regions of interest in our study. We appreciate your concern and hope this clarification provides sufficient rationale for our sequence parameters.

We have made this clearer in the revised manuscript.

Methods, Page 20-21 Lines 398-408

“Then, the “sleep” session began after the participants were instructed to try and fall asleep. For the functional scans, whole-brain images were acquired using k-space and steady-state T2*-weighted gradient echo-planar imaging (EPI) sequence that is sensitive to the BOLD contrast. This measures local magnetic changes caused by changes in blood oxygenation that accompany neural activity (sequence specification: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm). A relatively long TR and larger voxel size were chosen to optimize SNR and reduce susceptibility-related distortions, which are critical in EEG-fMRI sleep studies where head motion and physiological noise can be substantial. The longer TR allowed whole-brain coverage with sufficient temporal resolution, while the larger voxel size helped enhance BOLD sensitivity and minimize partial volume effects in deep brain structures (e.g., the thalamus and hippocampus), which are key regions of interest in this study.”

(8) The anatomically defined ROIs are quite large. It should be elaborated on how this might reduce sensitivity to sleep rhythm-specific activity within sub-regions, especially for the thalamus, which has distinct nuclei involved in sleep functions.

We appreciate your insight regarding the use of anatomically defined ROIs and their potential limitations in detecting sleep rhythm-specific activity within sub-regions, particularly in the thalamus. Given the distinct functional roles of thalamic nuclei in sleep processes, we acknowledge that using a single, large thalamic ROI may reduce sensitivity to localized activity patterns. To address this, we will discuss this limitation in the revised manuscript, acknowledging that our approach prioritizes whole-structure effects but may not fully capture nucleus-specific contributions.

Discussion, Page 18 Lines 333-341

“Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.”

(9) The study reports SO & spindle amplitudes & densities, as well as SO+spindle coupling, to be larger during N2/3 sleep compared to N1 and REM sleep, which is trivial but can be seen as a sanity check of the data. However, the amount of SOs and spindles reported for N1 and REM sleep is concerning, as per definition there should be hardly any (if SOs or spindles occur in N1 it becomes by definition N2, and the interval between spindles has to be considerably large in REM to still be scored as such). Thus, on the one hand, the report of these comparisons takes too much space in the main manuscript as it is trivial, but on the other hand, it raises concerns about the validity of the scoring.

We appreciate your concern regarding the reported presence of SOs and spindles in N1 and REM sleep and the potential implications. Our detection method for detecting SO, spindle, and coupling were originally designed only for N2&N3 sleep data based on the characteristics of the data itself, and this method is widely recognized and used in the sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). While, because the detection methods for SO and spindle are based on percentiles, this method will always detect a certain number of events when used for other stages (N1 and REM) sleep data, but the differences between these events and those detected in stage N23 remain unclear. We will acknowledge the reasons for these results in the Methods section and emphasize that they are used only for sanity checks.

Methods, Page 25 Lines 515-524

“We note that the above methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM (see also Figure S2-S4, Table S1-S4).”

(10) Why was electrode F3 used to quantify the occurrence of SOs and spindles? Why not a midline frontal electrode like Fz (or a number of frontal electrodes for SOs) and Cz (or a number of centroparietal electrodes) for spindles to be closer to their maximum topography?

We appreciate your suggestion regarding electrode selection for SO and spindle quantification. Our choice of F3 was primarily based on previous studies (Massimini et al., 2004; Molle et al., 2011), where bilateral frontal electrodes are commonly used for detecting SOs and spindles. Additionally, we considered the impact of MRI-related noise and, after a comprehensive evaluation, determined that F3 provided an optimal balance between signal quality and artifact minimization. We also acknowledge that alternative electrode choices, such as Fz for SOs and Cz for spindles, could provide additional insights into their topographical distributions.

(11) Functional connectivity (hippocampus -> thalamus -> cortex (mPFC)) is reported to be increased during SO-spindle coupling and interpreted as evidence for coordination of hippocampo-neocortical communication likely by thalamic spindles. However, functional connectivity was only analysed during coupled SO+spindle events, not during isolated SOs or isolated spindles. Without the direct comparison of the connectivity patterns between these three events, it remains unclear whether this is specific for coupled SO+spindle events or rather associated with one or both of the other isolated events. The PPIs need to be conducted for those isolated events as well and compared statistically to the coupled events.

We appreciate your critical perspective on our functional connectivity analysis and the interpretation of hippocampus-thalamus-cortex (mPFC) interactions during SO-spindle coupling. We acknowledge that, in the current analysis, functional connectivity was only examined during coupled SO-spindle events, without direct comparison to isolated SOs or isolated spindles. To address this concern, we have conducted PPI analyses for all three ROIs(Hippocampus, Thalamus, mPFC) and all three event types (SO-spindle couplings, isolated SOs, and isolated spindles). Our results indicate that neither isolated SOs nor isolated Spindles yielded significant connectivity changes in all three ROIs, as all failed to survive multiple comparison corrections. This suggests that the observed connectivity increase is specific to SO-spindle coupling, rather than being independently driven by either SOs or spindles alone.

Results, Page 14 Lines 248-255

“Crucially, the interaction between FC and SO-spindle coupling revealed that only the functional connectivity of hippocampus -> thalamus (ROI analysis, t(106) = 1.86, p = 0.0328) and thalamus -> mPFC (ROI analysis, t(106) = 1.98, p = 0.0251) significantly increased during SO-spindle coupling, with no significant changes in all other pathways (Fig. 4e). We also conducted PPI analyses for the other two events (SOs and spindles), and neither yielded significant connectivity changes in the three ROIs, as all failed to survive whole-brain FWE correction at the cluster level (p < 0.05). Together, these findings suggest that the thalamus, likely via spindles, coordinates hippocampal-cortical communication selectively during SO-spindle coupling, but not isolated SOs or spindle events alone.”

(12) The limited temporal resolution of fMRI does indeed not allow for easily distinguishing between fMRI activation patterns related to SO-up- vs. SO-down-states. For this, one could try to extract the amplitudes of SO-up- and SO-down-states separately for each SO event and model them as two separate parametric modulators (with the risk of collinearity as they are likely correlated).

We appreciate your insightful comment regarding the challenge of distinguishing fMRI activation patterns related to SO-up vs. SO-down states due to the limited temporal resolution of fMRI. While our current analysis does not differentiate between these two phases, we acknowledge that separately modeling SO-up and SO-down states using parametric modulators could provide a more refined understanding of their distinct neural correlates. However, as you notes, this approach carries the risk of collinearity, and there is indeed a high correlation between the two amplitudes across all subjects in our results (r=0.98). Future studies could explore more on leveraging high-temporal-resolution techniques. While implementing this in the current study is beyond our scope, we will acknowledge this limitation in the Discussion section.

Discussion, Page 17 Lines 308-322

“An intriguing aspect of our findings is the reduced DMN activity during SOs when modeled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.”

Discussion, Page 18 Lines 333-341

“Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.”

(13) L327: "It is likely that our findings of diminished DMN activity reflect brain activity during the SO DOWN-state, as this state consistently shows higher amplitude compared to the UP-state within subjects, which is why we modelled the SO trough as its onset in the fMRI analysis." This conclusion is not justified as the fact that SO down-states are larger in amplitude does not mean their impact on the BOLD response is larger.

We appreciate your concern regarding our interpretation of diminished DMN activity reflecting the SO down-state. We acknowledge that the current expression is somewhat misleading, and our interpretation of it is: it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. And we will make this clear in the Discussion section.

Discussion, Page 17 Lines 308-322

“An intriguing aspect of our findings is the reduced DMN activity during SOs when modeled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.”

(14) Line 77: "In the current study, while directly capturing hippocampal ripples with scalp EEG or fMRI is difficult, we expect to observe hippocampal activation in fMRI whenever SOs-spindles coupling is detected by EEG, if SOs- spindles-ripples triple coupling occurs during human NREM sleep". Not all SO-spindle events are associated with ripples (Staresina et al., 2015), but hippocampal activation may also be expected based on the occurrence of spindles alone (Bergmann et al., 2012).

We appreciate your clarification regarding the relationship between SO-spindle coupling and hippocampal ripples. We acknowledge that not all SO-spindle events are necessarily accompanied by ripples (Staresina et al., 2015). However, based on previous research, we found that hippocampal ripples are significantly more likely to occur during SO-spindle coupling events. This suggests that while ripple occurrence is not guaranteed, SO-spindle coupling creates a favorable network state for ripple generation and potential hippocampal activation. To ensure accuracy, we will revise the manuscript to delete this misleading sentence in the Introduction section and acknowledge in the Discussion that our results cannot conclusively directly observe the triple coupling of SO, spindle, and hippocampal ripples.

Discussion, Page 18 Lines 333-341

“Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.”

Reviewer #2 (Public review):

In this study, Wang and colleagues aimed to explore brain-wide activation patterns associated with NREM sleep oscillations, including slow oscillations (SOs), spindles, and SO-spindle coupling events. Their findings reveal that SO-spindle events corresponded with increased activation in both the thalamus and hippocampus. Additionally, they observed that SO-spindle coupling was linked to heightened functional connectivity from the hippocampus to the thalamus, and from the thalamus to the medial prefrontal cortex-three key regions involved in memory consolidation and episodic memory processes.

This study's findings are timely and highly relevant to the field. The authors' extensive data collection, involving 107 participants sleeping in an fMRI while undergoing simultaneous EEG recording, deserves special recognition. If shared, this unique dataset could lead to further valuable insights. While the conclusions of the data seem overall well supported by the data, some aspects with regard to the detection of sleep oscillations need clarification.

The authors report that coupled SO-spindle events were most frequent during NREM sleep (2.46 [plus minus] 0.06 events/min), but they also observed a surprisingly high occurrence of these events during N1 and REM sleep (2.23 [plus minus] 0.09 and 2.32 [plus minus] 0.09 events/min, respectively), where SO-spindle coupling would not typically be expected. Combined with the relatively modest SO amplitudes reported (~25 µV, whereas >75 µV would be expected when using mastoids as reference electrodes), this raises the possibility that the parameters used for event detection may not have been conservative enough - or that sleep staging was inaccurately performed. This issue could present a significant challenge, as the fMRI findings are largely dependent on the reliability of these detected events.

Thank you very much for your thorough and encouraging review. We appreciate your recognition of the significance and relevance of our study and dataset, particularly in highlighting how simultaneous EEG-fMRI recordings can provide complementary insights into the temporal dynamics of neural oscillations and their associated spatial activation patterns during sleep. In the sections that follow, we address each of your comments in detail. We have revised the text and conducted additional analyses wherever possible to strengthen our argument, clarify our methodological choices. We believe these revisions improve the clarity and rigor of our work, and we thank you for helping us refine it.

We appreciate your insightful comments regarding the detection of sleep oscillations. Our methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM. We will acknowledge the reasons for these results in the Methods section and emphasize that they are used only for sanity checks.

Regarding the reported SO amplitudes (~25 µV), during preprocessing, we applied the Signal Space Projection (SSP) method to more effectively remove MRI gradient artifacts and cardiac pulse noise. While this approach enhances data quality, it also reduces overall signal power, leading to systematically lower reported amplitudes. Despite this, our SO detection in NREM sleep (especially N2/N3) remain physiologically meaningful and are consistent with previous fMRI studies using similar artifact removal techniques. We appreciate your careful evaluation and valuable suggestions.

In addition, we will provide comprehensive tables in the supplementary materials, contains descriptive information about sleep-related characteristics (Table S1), as well as detailed information about sleep waves at each sleep stage for all 107 subjects(Table S2-S4), listing for each subject:(1)Different sleep stage duration; (2)Number of detected SOs; (3)Number of detected spindles; (4)Number of detected SO-spindle coupling events; (2)Density of detected SOs; (3)Density of detected spindles; (4)Density of detected SO-spindle coupling events.

Methods, Page 25 Lines 515-524

“We note that the above methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM (see also Figure S2-S4, Table S1-S4).”

Supplementary Materials, Page 42-54, Table S1-S4

Reviewer #3 (Public review):

Summary:

Wang et al., examined the brain activity patterns during sleep, especially when locked to those canonical sleep rhythms such as SO, spindle, and their coupling. Analyzing data from a large sample, the authors found significant coupling between spindles and SOs, particularly during the upstate of the SO. Moreover, the authors examined the patterns of whole-brain activity locked to these sleep rhythms. To understand the functional significance of these brain activities, the authors further conducted open-ended cognitive state decoding and found a variety of cognitive processing may be involved during SO-spindle coupling and during other sleep events. The authors next investigated the functional connectivity analyses and found enhanced connectivity between the hippocampus, the thalamus, and the medial PFC. These results reinforced the theoretical model of sleep-dependent memory consolidation, such that SO-spindle coupling is conducive to systems-level memory reactivation and consolidation.

Strengths:

There are obvious strengths in this work, including the large sample size, state-of-the-art neuroimaging and neural oscillation analyses, and the richness of results.

Weaknesses:

Despite these strengths and the insights gained, there are weaknesses in the design, the analyses, and inferences.

Thank you for your detailed and thoughtful review of our manuscript. We are delighted that you recognize our advanced analysis methods and rich results of neuroimaging and neural oscillations as well as the large sample size data. In the following sections, we provide detailed responses to each of your comments. And we have revised the text and conducted additional analyses to strengthen our arguments and clarify our methodological choices. We believe these revisions enhance the clarity and rigor of our work, and we sincerely appreciate your thoughtful feedback in helping us refine the manuscript.

(1) A repeating statement in the manuscript is that brain activity could indicate memory reactivation and thus consolidation. This is indeed a highly relevant question that could be informed by the current data/results. However, an inherent weakness of the design is that there is no memory task before and after sleep. Thus, it is difficult (if not impossible) to make a strong argument linking SO/spindle/coupling-locked brain activity with memory reactivation or consolidation.

We appreciate your suggestion regarding the lack of a pre- and post-sleep memory task in our study design. We acknowledge that, in the absence of behavioral measures, it is hard to directly link SO-spindle coupling to memory consolidation in an outcome-driven manner. Our interpretation is instead based on the well-established role of these oscillations in memory processes, as demonstrated in previous studies. We sincerely appreciate this feedback and will adjust our Discussion accordingly to reflect a more precise interpretation of our findings.

Discussion, Page 18 Lines 333-341

“Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.”

(2) Relatedly, to understand the functional implications of the sleep rhythm-locked brain activity, the authors employed the "open-ended cognitive state decoding" method. While this method is interesting, it is rather indirect given that there were no behavioral indices in the manuscript. Thus, discussions based on these analyses are speculative at best. Please either tone down the language or find additional evidence to support these claims.

Moreover, the results from this method are difficult to understand. Figure 3e showed that for all three types of sleep events (SO, spindle, SO-spindle), the same mental states (e.g., working memory, episodic memory, declarative memory) showed opposite directions of activation (left and right panels showed negative and positive activation, respectively). How to interpret these conflicting results? This ambiguity is also reflected by the term used: declarative memory and episodic memories are both indexed in the results. Yet these two processes can be largely overlapped. So which specific memory processes do these brain activity patterns reflect? The Discussion shall discuss these results and the limitations of this method.

We appreciate your critical assessment of the open-ended cognitive state decoding method and its interpretational challenges. Given the concerns about the indirectness of this approach, we decided to remove its related content and results from Figure 3 in the main text and include it in Supplementary Figure 7.

Due to the complexity of memory-related processes, we acknowledge that distinguishing between episodic and declarative memory based solely on this approach is not straightforward. We will revise the Supplementary Materials to explicitly discuss these limitations and clarify that our findings do not isolate specific cognitive processes but rather suggest general associations with memory-related networks.

Discussion, Page 17-18 Lines 323-332

“To explore functional relevance, we employed an open-ended cognitive state decoding approach using meta-analytic data (NeuroSynth: Yarkoni et al. (2011)). Although this method usefully generates hypotheses about potential cognitive processes, particularly in the absence of a pre- and post-sleep memory task, it is inherently indirect. Many cognitive terms showed significant associations (16 of 50), such as “episodic memory,” “declarative memory,” and “working memory.” We focused on episodic/declarative memory given the known link with hippocampal reactivation (Diekelmann & Born, 2010; Staresina et al., 2015; Staresina et al., 2023). Nonetheless, these inferences regarding memory reactivation should be interpreted cautiously without direct behavioral measures. Future research incorporating explicit tasks before and after sleep would more rigorously validate these potenial functional claims.”

(3) The coupling strength is somehow inconsistent with prior results (Hahn et al., 2020, eLife, Helfrich et al., 2018, Neuron). Specifically, Helfrich et al. showed that among young adults, the spindle is coupled to the peak of the SO. Here, the authors reported that the spindles were coupled to down-to-up transitions of SO and before the SO peak. It is possible that participants' age may influence the coupling (see Helfrich et al., 2018). Please discuss the findings in the context of previous research on SO-spindle coupling.

We appreciate your concern regarding the temporal characteristics of SO-spindle coupling. We acknowledge that the SO-spindle coupling phase results in our study are not identical to those reported by Hahn et al. (2020); Helfrich et al. (2018). However, these differences may arise due to slight variations in event detection parameters, which can influence the precise phase estimation of coupling. Notably, Hahn et al. (2020) also reported slight discrepancies in their group-level coupling phase results, highlighting that methodological differences can contribute to variability across studies. Furthermore, our findings are consistent with those of Schreiner et al. (2021), further supporting the robustness of our observations.

That said, we acknowledge that our original description of SO-spindle coupling as occurring at the "transition from the lower state to the upper state" was not entirely precise. The -π/2 phase represents the true transition point, while our observed coupling phase is actually closer to the SO peak rather than strictly at the transition. We will revise this statement in the manuscript to ensure clarity and accuracy in describing the coupling phase.

Discussion, Page 16 Lines 283-291

“Our data provide insights into the neurobiological underpinnings of these sleep rhythms. SOs, originating mainly in neocortical areas such as the mPFC, alternate between DOWN- and UP-states. The thalamus generates sleep spindles, which in turn couple with SOs. Our finding that spindle peaks consistently occurred slightly before the UP-state peak of SOs (in 83 out of 107 participants), concurs with prior studies, including Schreiner et al. (2021). Yet it differs from some results suggesting spindles might peak right at the SO UP-state (Hahn et al., 2020; Helfrich et al., 2018). Such discrepancies could arise from differences in detection algorithms, participant age (Helfrich et al., 2018), or subtle variations in cortical-thalamic timing. Nonetheless, these results underscore the importance of coordinated SO-spindle interplay in supporting sleep-dependent processes.”

(4) The discussion is rather superficial with only two pages, without delving into many important arguments regarding the possible functional significance of these results. For example, the author wrote, "This internal processing contrasts with the brain patterns associated with external tasks, such as working memory." Without any references to working memory, and without delineating why WM is considered as an external task even working memory operations can be internal. Similarly, for the interesting results on SO and reduced DMN activity, the authors wrote "The DMN is typically active during wakeful rest and is associated with self-referential processes like mind-wandering, daydreaming, and task representation (Yeshurun, Nguyen, & Hasson, 2021). Its reduced activity during SOs may signal a shift towards endogenous processes such as memory consolidation." This argument is flawed. DMN is active during self-referential processing and mind-wandering, i.e., when the brain shifts from external stimuli processing to internal mental processing. During sleep, endogenous memory reactivation and consolidation are also part of the internal mental processing given the lack of external environmental stimulation. So why during SO or during memory consolidation, the DMN activity would be reduced? Were there differences in DMN activity between SO and SO-spindle coupling events?

We appreciate your concerns regarding the brevity of the discussion and the need for clearer theoretical arguments. We will expand this section to provide more in-depth interpretations of our findings in the context of prior literature. Regarding working memory (WM), we acknowledge that our phrasing was ambiguous. We will modify this statement in the Discussion section.

For the SO-related reduction in DMN activity, we recognize the need for a more precise explanation. This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state.

To address your final question, we have conducted the additional post hoc comparison of DMN activity between isolated SOs and SO-spindle coupling events. Our results indicate that

DMN activation during SOs was significantly lower than during SO-spindle coupling (t(106) = -4.17, p < 1e-4). This suggests that SO-spindle coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. We appreciate your constructive feedback and will integrate these expanded analyses and discussions into our revised manuscript.

Results, Page 11 Lines 199-208

“Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). However, no decrease in DMN activity was found during SO-spindle coupling, and DMN activity during SO was significantly lower than during coupling (ROI analysis, t(106) = -4.17, p < 1e-4). For more detailed activation patterns, see Table S5-S7. We also varied the threshold used to detect SO events to assess its effect on hippocampal activation during SO-spindle coupling and observed that hippocampal activation remained significant when the percentile thresholds for SO detection ranged between 71% and 80% (see Fig. S6).”

Discussion, Page 17-18 Lines 308-332

“An intriguing aspect of our findings is the reduced DMN activity during SOs when modeled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.

To explore functional relevance, we employed an open-ended cognitive state decoding approach using meta-analytic data (NeuroSynth: Yarkoni et al. (2011)). Although this method usefully generates hypotheses about potential cognitive processes, particularly in the absence of a pre- and post-sleep memory task, it is inherently indirect. Many cognitive terms showed significant associations (16 of 50), such as “episodic memory,” “declarative memory,” and “working memory.” We focused on episodic/declarative memory given the known link with hippocampal reactivation (Diekelmann & Born, 2010; Staresina et al., 2015; Staresina et al., 2023). Nonetheless, these inferences regarding memory reactivation should be interpreted cautiously without direct behavioral measures. Future research incorporating explicit tasks before and after sleep would more rigorously validate these potential functional claims.”

Recommendations for the authors:

Reviewing Editor Comment:

The reviewers think that you are working on a relevant and important topic. They are praising the large sample size used in the study. The reviewers are not all in line regarding the overall significance of the findings, but they all agree the paper would strongly benefit from some extra work, as all reviewers raise various critical points that need serious consideration.

We appreciate your recognition of the relevance and importance of our study, as well as your acknowledgment of the large sample size as a strength of our work. We understand that there are differing perspectives regarding the overall significance of our findings, and we value the constructive critiques provided. We are committed to addressing the key concerns raised by all reviewers, including refining our analyses, clarifying our interpretations, and incorporating additional discussions to strengthen the manuscript. Below, we address your specific recommendations and provide responses to each point you raised to ensure our methods and results are as transparent and comprehensible as possible. We believe that these revisions will significantly enhance the rigor and impact of our study, and we sincerely appreciate your thoughtful feedback in helping us improve our work.

Reviewer #1 (Recommendations for the authors):

(1) The phrase "overnight sleep" suggests an entire night, while these were rather "nocturnal naps". Please rephrase.

Response: Thank you for pointing this out. We have revised the phrasing in our manuscript to "nocturnal naps" instead of "overnight sleep" to more accurately reflect the duration of the sleep recordings.

(2) Sleep staging results (macroscopic sleep architecture) should be provided in more detail (at least min and % of the different sleep stages, sleep onset latency, total sleep duration, total recording duration), at least mean/SD/range.

Thank you for this suggestion. We will provide comprehensive tables in the supplementary materials, contains descriptive information about sleep-related characteristics. This information will help provide a clearer overview of the macroscopic sleep architecture in our dataset.

Reviewer #2 (Recommendations for the authors):

In order to allow for a better estimation of the reliability of the detected sleep events, please:

(1) Provide densities and absolute numbers of all detected SOs and spindles (N1, NREM, and REM sleep).

Thank you for pointing this out. We will provide comprehensive tables in the supplementary materials, contains detailed information about sleep waves at each sleep stage for all 107 subjects (Table S2-S4), listing for each subject:1) Different sleep stage duration; 2) Number of detected SOs; 3) Number of detected spindles; 4) Number of detected SO-spindle coupling events; 5) Density of detected SOs; 6) Density of detected spindles; 7) Density of detected SO-spindle coupling events.

Supplementary Materials, Page 43-54, Table S2-S4

(2) Show ERPs for all detected SOs and spindles (per sleep stage).

Thank you for the suggestion. We will provide ERPs for all detected SOs and spindles, separated by sleep stage (N1, N2&N3, and REM) in supplementary Fig. S2-S4. These ERP waveforms will help illustrate the characteristic temporal profiles of SOs and spindles across different sleep stages.

Methods, Page 25, Line 525-532

“Event-related potentials (ERP) analysis. After completing the detection of each sleep rhythm event, we performed ERP analyses for SOs, spindles, and coupling events in different sleep stages. Specifically, for SO events, we took the trough of the DOWN-state of each SO as the zero-time point, then extracted data in a [-2 s to 2 s] window from the broadband (0.1–30 Hz) EEG and used [-2 s to -0.5 s] for baseline correction; the results were then averaged across 107 subjects (see Fig. S2a). For spindle events, we used the peak of each spindle as the zero-time point and applied the same data extraction window and baseline correction before averaging across 107 subjects (see Fig. S2b). Finally, for SO-spindle coupling events, we followed the same procedure used for SO events (see Fig. 2a, Figs. S3–S4).”

(3) Provide detailed info concerning sleep characteristics (time spent in each sleep stage etc.).

Thank you for this suggestion. Same as the response above, we will provide comprehensive tables in the supplementary materials, contains descriptive information about sleep-related characteristics.

Supplementary Materials, Page 42, Table S1 (same as above)

(4) What would happen if more stringent parameters were used for event detection? Would the authors still observe a significant number of SO spindles during N1 and REM? Would this affect the fMRI-related results?

Thank you for this suggestion. Our methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM (see also Figure S2-S4, Table S1-S4).

Furthermore, in order to explore the impact of this on our fMRI results, we conducted an additional sensitivity analysis by applying different detection parameters for SOs. Specifically, we adjusted amplitude percentile thresholds for SO detection (the parameter that has the greatest impact on the results). We used the hippocampal activation value during N2&N3 stage SO-spindle coupling as an anchor value and found that when the parameters gradually became stricter, the results were similar to or even better than the current results. However, when we continued to increase the threshold, the results began to gradually decrease until the threshold was increased to 80%, and the results were no longer significant. This indicates that our results are robust within a specific range of parameters, but as the threshold increases, the number of trials decreases, ultimately weakening the statistical power of the fMRI analysis.

Thank you again for your suggestions on sleep rhythm event detection. We will add the results in Supplementary and revise our manuscript accordingly.

Results, Page 11, Line 199-208

“Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). However, no decrease in DMN activity was found during SO-spindle coupling, and DMN activity during SO was significantly lower than during coupling (ROI analysis, t(106) = -4.17, p < 1e-4). For more detailed activation patterns, see Table S5-S7. We also varied the threshold used to detect SO events to assess its effect on hippocampal activation during SO-spindle coupling and observed that hippocampal activation remained significant when the percentile thresholds for SO detection ranged between 71% and 80% (see Fig. S6).”

Finally, we sincerely thank all again for your thoughtful and constructive feedback. Your insights have been invaluable in refining our analyses, strengthening our interpretations, and improving the clarity and rigor of our manuscript. We appreciate the time and effort you have dedicated to reviewing our work, and we are grateful for the opportunity to enhance our study based on your recommendations.

References:

Bergmann, T. O., Mölle, M., Diedrichs, J., Born, J., & Siebner, H. R. (2012). Sleep spindle-related reactivation of category-specific cortical regions after learning face-scene associations. NeuroImage, 59(3), 2733-2742.

Buzsáki, G. (2015). Hippocampal sharp wave‐ripple: A cognitive biomarker for episodic memory and planning. Hippocampus, 25(10), 1073-1188.

Caporro, M., Haneef, Z., Yeh, H. J., Lenartowicz, A., Buttinelli, C., Parvizi, J., & Stern, J. M. (2012). Functional MRI of sleep spindles and K-complexes. Clinical neurophysiology, 123(2), 303-309.

Coulon, P., Budde, T., & Pape, H.-C. (2012). The sleep relay—the role of the thalamus in central and decentral sleep regulation. Pflügers Archiv-European Journal of Physiology, 463, 53-71.

Crunelli, V., Lőrincz, M. L., Connelly, W. M., David, F., Hughes, S. W., Lambert, R. C., Leresche, N., & Errington, A. C. (2018). Dual function of thalamic low-vigilance state oscillations: rhythm-regulation and plasticity. Nature Reviews Neuroscience, 19(2), 107-118.

Czisch, M., Wehrle, R., Stiegler, A., Peters, H., Andrade, K., Holsboer, F., & Sämann, P. G. (2009). Acoustic oddball during NREM sleep: a combined EEG/fMRI study. PloS one, 4(8), e6749.

Diba, K., & Buzsáki, G. (2007). Forward and reverse hippocampal place-cell sequences during ripples. Nature Neuroscience, 10(10), 1241.

Diekelmann, S., & Born, J. (2010). The memory function of sleep. Nature Reviews Neuroscience, 11(2), 114-126.

Fogel, S., Albouy, G., King, B. R., Lungu, O., Vien, C., Bore, A., Pinsard, B., Benali, H., Carrier, J., & Doyon, J. (2017). Reactivation or transformation? Motor memory consolidation associated with cerebral activation time-locked to sleep spindles. PloS one, 12(4), e0174755.

Hahn, M. A., Heib, D., Schabus, M., Hoedlmoser, K., & Helfrich, R. F. (2020). Slow oscillation-spindle coupling predicts enhanced memory formation from childhood to adolescence. Elife, 9, e53730.

Halassa, M. M., Siegle, J. H., Ritt, J. T., Ting, J. T., Feng, G., & Moore, C. I. (2011). Selective optical drive of thalamic reticular nucleus generates thalamic bursts and cortical spindles. Nature Neuroscience, 14(9), 1118-1120.

Hale, J. R., White, T. P., Mayhew, S. D., Wilson, R. S., Rollings, D. T., Khalsa, S., Arvanitis, T. N., & Bagshaw, A. P. (2016). Altered thalamocortical and intra-thalamic functional connectivity during light sleep compared with wake. NeuroImage, 125, 657-667.

Helfrich, R. F., Lendner, J. D., Mander, B. A., Guillen, H., Paff, M., Mnatsakanyan, L., Vadera, S., Walker, M. P., Lin, J. J., & Knight, R. T. (2019). Bidirectional prefrontal-hippocampal dynamics organize information transfer during sleep in humans. Nature Communications, 10(1), 3572.

Helfrich, R. F., Mander, B. A., Jagust, W. J., Knight, R. T., & Walker, M. P. (2018). Old brains come uncoupled in sleep: slow wave-spindle synchrony, brain atrophy, and forgetting. Neuron, 97(1), 221-230. e224.

Horovitz, S. G., Fukunaga, M., de Zwart, J. A., van Gelderen, P., Fulton, S. C., Balkin, T. J., & Duyn, J. H. (2008). Low frequency BOLD fluctuations during resting wakefulness and light sleep: A simultaneous EEG‐fMRI study. Human brain mapping, 29(6), 671-682.

Huang, Q., Xiao, Z., Yu, Q., Luo, Y., Xu, J., Qu, Y., Dolan, R., Behrens, T., & Liu, Y. (2024). Replay-triggered brain-wide activation in humans. Nature Communications, 15(1), 7185.

Ilhan-Bayrakcı, M., Cabral-Calderin, Y., Bergmann, T. O., Tüscher, O., & Stroh, A. (2022). Individual slow wave events give rise to macroscopic fMRI signatures and drive the strength of the BOLD signal in human resting-state EEG-fMRI recordings. Cerebral Cortex, 32(21), 4782-4796.

Laufs, H. (2008). Endogenous brain oscillations and related networks detected by surface EEG‐combined fMRI. Human brain mapping, 29(7), 762-769.

Laufs, H., Walker, M. C., & Lund, T. E. (2007). ‘Brain activation and hypothalamic functional connectivity during human non-rapid eye movement sleep: an EEG/fMRI study’—its limitations and an alternative approach. Brain, 130(7), e75.

Margulies, D. S., Ghosh, S. S., Goulas, A., Falkiewicz, M., Huntenburg, J. M., Langs, G., Bezgin, G., Eickhoff, S. B., Castellanos, F. X., & Petrides, M. (2016). Situating the default-mode network along a principal gradient of macroscale cortical organization. Proceedings of the National Academy of Sciences, 113(44), 12574-12579.

Massimini, M., Huber, R., Ferrarelli, F., Hill, S., & Tononi, G. (2004). The sleep slow oscillation as a traveling wave. Journal of Neuroscience, 24(31), 6862-6870.

Moehlman, T. M., de Zwart, J. A., Chappel-Farley, M. G., Liu, X., McClain, I. B., Chang, C., Mandelkow, H., Özbay, P. S., Johnson, N. L., & Bieber, R. E. (2019). All-night functional magnetic resonance imaging sleep studies. Journal of neuroscience methods, 316, 83-98.

Molle, M., Bergmann, T. O., Marshall, L., & Born, J. (2011). Fast and slow spindles during the sleep slow oscillation: disparate coalescence and engagement in memory processing. Sleep, 34(10), 1411-1421.

Ngo, H.-V., Fell, J., & Staresina, B. (2020). Sleep spindles mediate hippocampal-neocortical coupling during long-duration ripples. Elife, 9, e57011.

Picchioni, D., Horovitz, S. G., Fukunaga, M., Carr, W. S., Meltzer, J. A., Balkin, T. J., Duyn, J. H., & Braun, A. R. (2011). Infraslow EEG oscillations organize large-scale cortical– subcortical interactions during sleep: a combined EEG/fMRI study. Brain research, 1374, 63-72.

Schabus, M., Dang-Vu, T. T., Albouy, G., Balteau, E., Boly, M., Carrier, J., Darsaud, A., Degueldre, C., Desseilles, M., & Gais, S. (2007). Hemodynamic cerebral correlates of sleep spindles during human non-rapid eye movement sleep. Proceedings of the National Academy of Sciences, 104(32), 13164-13169.

Schreiner, T., Kaufmann, E., Noachtar, S., Mehrkens, J.-H., & Staudigl, T. (2022). The human thalamus orchestrates neocortical oscillations during NREM sleep. Nature communications, 13(1), 5231.

Schreiner, T., Petzka, M., Staudigl, T., & Staresina, B. P. (2021). Endogenous memory reactivation during sleep in humans is clocked by slow oscillation-spindle complexes. Nature Communications, 12(1), 3112.

Singh, D., Norman, K. A., & Schapiro, A. C. (2022). A model of autonomous interactions between hippocampus and neocortex driving sleep-dependent memory consolidation. Proceedings of the National Academy of Sciences, 119(44), e2123432119.

Spoormaker, V. I., Schröter, M. S., Gleiser, P. M., Andrade, K. C., Dresler, M., Wehrle, R., Sämann, P. G., & Czisch, M. (2010). Development of a large-scale functional brain network during human non-rapid eye movement sleep. Journal of Neuroscience, 30(34), 11379-11387.

Staresina, B. P., Bergmann, T. O., Bonnefond, M., van der Meij, R., Jensen, O., Deuker, L., Elger, C. E., Axmacher, N., & Fell, J. (2015). Hierarchical nesting of slow oscillations, spindles and ripples in the human hippocampus during sleep. Nature Neuroscience, 18(11), 1679-1686.

Staresina, B. P., Niediek, J., Borger, V., Surges, R., & Mormann, F. (2023). How coupled slow oscillations, spindles and ripples coordinate neuronal processing and communication during human sleep. Nature Neuroscience, 1-9.

Yarkoni, T., Poldrack, R. A., Nichols, T. E., Van Essen, D. C., & Wager, T. D. (2011). Large-scale automated synthesis of human functional neuroimaging data. Nature methods, 8(8), 665-670.

Yeshurun, Y., Nguyen, M., & Hasson, U. (2021). The default mode network: where the idiosyncratic self meets the shared social world. Nature Reviews Neuroscience, 1-12.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: During sleep, our brain undergoes highly synchronized activity, orchestrated by distinct neural rhythms. Little is known about the associated brain activation during these sleep rhythms, and even less about their functional implications. In this study, we investigated the brain-wide activation underlying human sleep rhythms by employing simultaneous electroencephalography (EEG) and functional magnetic resonance imaging (fMRI) in 107 participants during overnight sleep. We identified a significant coupling between slow oscillations (SO) and spindle events during non-rapid eye movement (NREM) sleep, particularly at the UP-state of SOs. This coupling was associated with increased activation in the thalamus and hippocampus, showing a brain-wide activation that resembles episodic memory processing, yet is distinctly dissociated from task-related activation. Moreover, this SO-spindle coupling was linked to a selective increase in functional connectivity from the hippocampus to the thalamus, and from the thalamus to the neocortex, particularly the medial prefrontal cortex. These findings suggest that the thalamus plays a crucial role in coordinating the hippocampal-cortical dialogue during sleep.

Revised version: During sleep, our brain undergoes highly synchronized activity, orchestrated by distinct neural rhythms. Little is known about the associated brain activation during these sleep rhythms, and even less about their functional implications. In this study, we investigated the brain-wide activation underlying human sleep rhythms by employing simultaneous electroencephalography (EEG) and functional magnetic resonance imaging (fMRI) in 107 participants during nocturnal nap (first half of the night). We identified robust coupling between slow oscillations (SOs) and fast spindles during deep non-rapid eye movement (NREM) sleep (N2/3 stages), with spindle peaks consistently occurring just before the SO UP-state. This SO-spindle coupling was linked to elevated activation in both the thalamus and hippocampus, alongside increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the medial prefrontal cortex (mPFC). An open-ended cognitive state decoding analysis suggested that these activations may relate to episodic memory processes, yet were distinct from task-related networks. Together, these findings highlight the thalamus as a key coordinator of hippocampal-cortical communication during sleep and provide new insights into the mechanisms by which synchronized sleep rhythms may support memory consolidation.

Introduction

Earlier version: During sleep, the human brain engages in highly synchronized activity driven by distinct rhythms. In non-rapid eye movement (NREM) sleep, the emergence and coordination of slow oscillations (SOs) and sleep spindles are key features (Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Staresina et al., 2015; Staresina et al., 2023). However, the brain activation patterns associated with these rhythms and their coordination remain largely unexplored, and their functional implications are even less understood.

Revised version: During sleep, the human brain engages in highly synchronized activity driven by distinct neural rhythms. In non-rapid eye movement (NREM) sleep, the emergence and coordination of slow oscillations (SOs) and sleep spindles are key features (Hahn et al., 2020; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Staresina et al., 2015; Staresina et al., 2023). However, the brain activation patterns associated with these rhythms, and especially their coordination, remain largely unexplored, and their functional implications are even less understood.

Earlier version: SOs (∼1 Hz) represent shifts in neuronal membrane potential between periods of silence (’hyperpolarization’, or ‘DOWN’ state) and excitation (’depolarization’,or ‘UP’ state) (Amzica & Steriade, 2002; Steriade, McCormick, & Sejnowski, 1993). These oscillations originate spontaneously in neocortical regions, particularly the medial prefrontal cortex (mPFC), during NREM sleep (Dang-Vu et al., 2008; Massimini et al., 2004; Nir et al., 2011). The depolarization phase of SOs is believed to trigger sleep spindles in the thalamus, producing oscillations in the 11–16 Hz range (Fernandez & L üthi, 2020; Mak-McCully et al., 2017). These spindles align with the excitable UP-states of SOs and synchronize with hippocampal ripples (Buzsáki, 2015; Helfrich et al., 2019; Joo & Frank, 2018; Ngo, Fell, & Staresina, 2020). Rodent studies using optogenetics have shown that only thalamic spindles phase-locked with the UP-state of cortical SOs enhance memory consolidation, while those out of phase do not (Latchoumane et al., 2017). This finding is consistent with intracranial evidence from human epilepsy patients, suggesting the thalamus acts as a key relay between the hippocampus and cortex (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Schreiner et al., 2022; Staresina et al., 2015).

Revised version: SOs (∼1 Hz) are characterized by fluctuations in neuronal membrane potential between periods of silence (’hyperpolarization’, or ‘DOWN’ state) and excitation (’depolarization’, or ‘UP’ state) (Amzica & Steriade, 2002; Steriade, McCormick, & Sejnowski, 1993). These oscillations originate in neocortical regions, particularly the medial prefrontal cortex (mPFC), during NREM sleep (Dang-Vu et al., 2008; Massimini et al., 2004; Nir et al., 2011). The depolarization phase of SOs is believed to trigger sleep spindles in the thalamus, producing 11–16 Hz oscillations (Fernandez & Lüthi, 2020; Mak-McCully et al., 2017). Thalamic spindles tend to align with the excitable UP-states of SOs and can synchronize with hippocampal ripples (Buzsáki, 2015; Helfrich et al., 2019; Joo & Frank, 2018; Ngo, Fell, & Staresina, 2020). Rodent optogenetics studies have shown that only thalamic spindles phase-locked to the UP-state of cortical SOs enhance memory consolidation, while out-of-phase spindles do not (Latchoumane et al., 2017). This finding is consistent with intracranial evidence from human epilepsy patients, suggesting the thalamus acts as a key relay between the hippocampus and cortex, coordinating their interaction (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Schreiner et al., 2022; Staresina et al., 2015).

Earlier version: The triple coupling of SOs, spindles, and ripples is believed to facilitate memory consolidation by synchronizing neuronal activity across brain regions and replaying memory traces from the hippocampus to the neocortex (Diekelmann & Born, 2010; Latchoumane et al., 2017; Liu et al., 2022; Singh, Norman, & Schapiro, 2022; Staresina et al., 2015; Staresina et al., 2023). However, the exact mechanisms of this inter-regional communication during coordinated sleep rhythms remain unclear.

Revised version: The triple coupling of SOs, spindles, and ripples is believed to facilitate memory consolidation by synchronizing neuronal activity across brain regions and replaying memory traces from the hippocampus to the neocortex (Diekelmann & Born, 2010; Latchoumane et al., 2017; Liu et al., 2022; Singh, Norman, & Schapiro, 2022; Staresina et al., 2015; Staresina et al., 2023), possibly through relay from thalamic spindles. However, the exact mechanisms of this inter-regional communication during coordinated sleep rhythms remain unclear. Understanding the brain-wide activation associated with these rhythms is crucial, as it reveals how information flows among brain regions, especially during the critical time windows when these rhythms align.

Added in the revised version: Investigating these sleep-related neural processes in humans is challenging because it requires tracking transient sleep rhythms while simultaneously assessing their widespread brain activation. Recent advances in simultaneous EEG-fMRI techniques provide a unique opportunity to explore these processes. EEG allows for precise event-based detection of neural signal, while fMRI provides insight into the broader spatial patterns of brain activation and functional connectivity (Horovitz et al., 2008; Huang et al., 2024; Laufs, 2008; Laufs, Walker, & Lund, 2007; Schabus et al., 2007; Spoormaker et al., 2010). Previous EEG-fMRI studies on sleep have focused on classifying sleep stages or examining the neural correlates of specific waves (Bergmann et al., 2012; Caporro et al., 2012; Czisch et al., 2009; Fogel et al., 2017; Hale et al., 2016; Ilhan-Bayrakcı et al., 2022; Moehlman et al., 2019; Picchioni et al., 2011). These studies have generally reported that slow oscillations are associated with widespread cortical and subcortical BOLD changes, whereas spindles elicit activation in the thalamus, as well as in several cortical and paralimbic regions. Although these findings provide valuable insights into the BOLD correlates of sleep rhythms, they often do not employ sophisticated temporal modeling (Huang et al., 2024), to capture the dynamic interactions between different oscillatory events, e.g., the coupling between SOs and spindles.

Removed in the revised version: Understanding the brain-wide activation associated with these rhythms is crucial, as it reveals how information flows among brain regions, especially during the critical time windows when these rhythms align.

Removed in the revised version: Studying these processes in the human brain is challenging because it requires tracking both the transient nature of sleep rhythms and their widespread brain activation simultaneously. Recent advances in simultaneous EEG-fMRI techniques, which allow researchers to model temporal events captured by EEG and explain the associated brain activity observed through fMRI, offer promising opportunities to capture these brain-wide activations (Huang et al., 2024). In the current study, while directly capturing hippocampal ripples with scalp EEG or fMRI is difficult, we expect to observe hippocampal activation in fMRI whenever SOs-spindles coupling is detected by EEG, if SOs-spindles-ripples triple coupling occurs during human NREM sleep.

Earlier version: Another advantage of identifying sleep rhythm-related brain activation is the potential to infer their potential cognitive functions through open-ended cognitive state decoding (Yarkoni et al., 2011). Since performing cognitive tasks during sleep is limited, most neuroimaging studies on human sleep rely on the ‘targeted memory reactivation (TMR)’ paradigm, where participants are re-exposed to sensory cues during sleep to trigger associated memory reactivation (Cousins et al., 2016; Hu et al., 2020; Oudiette & Paller, 2013; Siefert et al., 2024). For instance, fMRI studies have shown significant hippocampal activation in response to odour re-exposure during sleep (Rasch et al., 2007), while EEG studies suggest that this reactivation is more likely during the UP-state of SOs (Cousins et al., 2016; Göldi et al., 2019; Xia et al., 2023). However, this approach provides limited insight into the neural mechanisms underlying natural sleep rhythms.

Revised version: Understanding sleep rhythm-related brain activation is also essential for uncovering its potential cognitive functions. Since performing cognitive tasks during sleep is limited, most neuroimaging studies on human sleep rely on the targeted memory reactivation (TMR) paradigm, where participants are re-exposed to sensory cues during sleep to trigger associated memory reactivation (Cousins et al., 2016; Hu et al., 2020; Oudiette & Paller, 2013; Siefert et al., 2024). For instance, fMRI studies have shown significant hippocampal activation in response to odour re-exposure during sleep (Rasch et al., 2007), while EEG studies suggest that this reactivation is more likely during the UP-state of SOs (Cousins et al., 2016; Göldi et al., 2019; Xia et al., 2023). However, while TMR provides insight into externally induced memory processing during sleep, it offers limited understanding of the endogenous neural mechanisms underlying natural sleep rhythms.

Added in the revised version: To address this gap, our study investigates brain-wide activation and functional connectivity patterns associated with SO-spindle coupling, and employs a cognitive state decoding approach (Margulies et al., 2016; Yarkoni et al., 2011)—albeit indirectly—to infer potential cognitive functions. In the current study, we used simultaneous EEG-fMRI recordings during nocturnal naps (detailed sleep staging results are provided in the Methods and Table S1) in 107 participants. Although directly detecting hippocampal ripples using scalp EEG or fMRI is challenging, we expected that hippocampal activation in fMRI would coincide with SO-spindle coupling detected by EEG, given that SOs, spindles, and ripples frequently co-occur during NREM sleep. We also anticipated a critical role of the thalamus, particularly thalamic spindles, in coordinating hippocampal-cortical communication.

Added in the revised version: We found significant coupling between SOs and spindles during NREM sleep (N2/3), with spindle peaks occurring slightly before the SO peak. This coupling was associated with increased activation in both the thalamus and hippocampus, with functional connectivity patterns suggesting thalamic coordination of hippocampal-cortical communication. These findings highlight the key role of the thalamus in coordinating hippocampal-cortical interactions during human sleep and provide new insights into the neural mechanisms underlying sleep-dependent brain communication. A deeper understanding of these mechanisms may contribute to future neuromodulation approaches aimed at enhancing sleep-dependent cognitive function and treating sleep-related disorders.

Removed in the revised version: By using cognitive state decoding, we can infer the functional implications of natural sleep rhythms by comparing their brain activation patterns with known task-related patterns (Margulies et al., 2016; Yarkoni et al., 2011). This method offers a promising way to bridge the gap between observed brain activity during sleep and its functional significance, providing new insights into the relationship between sleep and cognition.

Removed in the revised version: In this study, we investigated brain-wide activation during sleep using simultaneous EEG and fMRI in 107 participants, combined with cognitive state decoding. Our findings revealed significant coupling between SOs and spindles during NREM sleep, particularly at the transition to UP-state of SOs. This coupling was associated with increased activation in the thalamus and hippocampus, and a brain-wide activation pattern resembling episodic memory processing, yet distinct from task-related activations. Furthermore, SO-spindle coupling was linked to increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC.

Removed in the revised version: These findings suggest that the thalamus plays a crucial role in coordinating hippocampal-cortical dialogue during human sleep, highlighting the importance of intrinsic brain-wide activation when sleep rhythms align. Understanding these mechanisms paves the way for developing neuromodulation interventions to enhance cognitive function during sleep and treat sleep-related disorders.

Results

Earlier version: Simultaneous EEG-fMRI data during nocturnal sleep were collected from 138 healthy subjects (age, 22.56 ± 0.30; 81 females). After excluding subjects due to excessive head motion (mean framewise displacement > 0.5) or insufficient scanning sleep duration (<10 min), 107 participants (age, 22.40 ± 0.34; 63 females) were retained for subsequent analysis, with an average scanning time of sleep of 3.01 ± 0.14 hours (see Methods). Since NREM sleep predominantly occurs in early sleep (Carskadon & Dement, 2005), we conducted our scans during the first half of the night to ensure sufficient coverage of NREM sleep, this is also the period when sleep rhythms are most consistently observed and studied in relation to memory consolidation (Maingret et al., 2016; Molle et al., 2011; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023).

Revised version: Simultaneous EEG-fMRI data during nocturnal sleep were collected from 138 healthy subjects (age, 22.56 ± 0.30; 81 females). After excluding subjects due to excessive head motion (mean framewise displacement > 0.5) or insufficient scanning sleep duration (<10 min), 107 participants (age, 22.40 ± 0.34; 63 females) were retained for subsequent analysis, with an average scanning time of sleep of 3.01 ± 0.14 hours (see Methods and Table S1). Since NREM sleep predominantly occurs in early sleep (Carskadon & Dement, 2005), we conducted our scans during the first half of the night to ensure sufficient coverage of NREM sleep, this is also the period when sleep rhythms are most consistently observed and studied in relation to memory consolidation (Maingret et al., 2016; Molle et al., 2011; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023).

Sleep stages and sleep rhythms

The primary objective of the EEG analysis was to examine the intrinsic characteristics of sleep rhythms and to identify the timing of these SO and spindle events for subsequent fMRI analyses.

Earlier version: We first removed MRI gradient noise from the EEG signal (detailed in Methods, see Fig. S1) and then applied an automated sleep staging algorithm (Vallat & Walker, 2021). The staging results were manually reviewed by two sleep experts to ensure the accurate classification of each sleep stages (Fig. 1b). Upon validation, we confirmed that N2/3 stages predominated in the dataset, with a significantly higher proportion of N2/3 sleep (75.77% ± 1.50%) compared to N1 (14.57% ± 1.06%, t(106) = 25.04, p < 1e-4, paired-samples t test) and REM sleep (9.66% ± 0.90%, t(106) = 29.60, p < 1e-4, Fig. 1d).

Revised version: We first removed MRI gradient noise from the EEG signal (detailed in Methods, see Fig. S1) and then applied an automated sleep staging algorithm (Vallat & Walker, 2021). The staging results were manually reviewed by two sleep experts to ensure the accurate classification of each sleep stages (Fig. 1b). Upon validation, we confirmed that N2/3 stages predominated in the dataset, with a significantly higher proportion of N2/3 sleep (75.77% ± 1.50%) compared to N1 (14.57% ± 1.06%, t(106) = 25.04, p < 1e-4, paired-samples t test) and REM sleep (9.66% ± 0.90%, t(106) = 29.60, p < 1e-4, Fig. 1d, Table S1).

Sleep stages and sleep rhythms in 107 subjects.

Earlier version: a, Sleep rhythms and task schematic. Subjects slept overnight with simultaneous EEG-fMRI recordings. Since detecting hippocampal ripples directly from scalp EEG is challenging, our focus was on capturing SOs, spindles, and their couplings. Regions of interest (ROIs) are colour-coded: green for the thalamus (spindle), purple for the mPFC (SOs), and orange for the hippocampus (ripples). b, Sleep staging and EEG spectrogram. N2/3 sleep stages (red line) were initially identified using an offline automatic sleep staging algorithm (Vallat & Walker, 2021) and then manually validated. (c) Schematic of EEG data across different sleep stages, using preprocessed data from the C3 electrode. (d) Proportion of each sleep stage in the dataset. (e) Amplitudes (μV) of SOs (left) and spindles (right) across sleep stages. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.

Revised version: a, Sleep rhythms and task schematic. Subjects slept first half of the night with simultaneous EEG-fMRI recordings. Since detecting hippocampal ripples directly from scalp EEG is challenging, our focus was on capturing SOs, spindles, and their couplings. Regions of interest (ROIs) are colour-coded: green for the thalamus (spindle), purple for the mPFC (SOs), and orange for the hippocampus (ripples). b, Sleep staging and EEG spectrogram. N2/3 sleep stages (red line) were initially identified using an offline automatic sleep staging algorithm (Vallat & Walker, 2021) and then manually validated. (c) Schematic of EEG data across different sleep stages, using preprocessed data from the C3 electrode. (d) Proportion of each sleep stage in the dataset. (e) Amplitudes (μV) of SOs (left) and spindles (right) across sleep stages. The SO amplitudes used here are from 0.16-1.25 Hz band EEG data, and the spindle amplitudes used here are from 12-16 Hz band EEG data. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.

Earlier version: Each sleep stage is characterized by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). As a sanity check, we compared SO and spindle amplitudes and densities across stages. Both were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV, spindle: 7.39 ± 0.27 μV) compared to N1 (SO: 18.87 ± 1.08 μV, spindle: 6.72 ± 0.25 μV) and REM (SO: 16.67 ± 1.14 μV, spindle: 6.23 ± 0.23 μV, all p < 1e-4, Fig. 1e). This aligns with the established understanding that N2/3 stages feature more pronounced slow-wave activity and robust spindles (Gaillard & Blois, 1981). Furthermore, we found that the amplitude of the SO DOWN-state (13.62 ± 0.67 μV) was significantly higher than the UP-state (12.98 ± 0.66 μV, t(106) = 4.82, p < 1e-4), consistent with previous work (Cash et al., 2009; Dang-Vu et al., 2008). Based on this, we used the DOWN-state to lock the timing of SOs for subsequent EEG-informed fMRI analysis.

Revised version: Each sleep stage is characterized by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). As a sanity check, we compared SO and spindle amplitudes and densities across stages. Both were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV, spindle: 7.39 ± 0.27 μV) compared to N1 (SO: 18.87 ± 1.08 μV, spindle: 6.72 ± 0.25 μV) and REM (SO: 16.67 ± 1.14 μV, spindle: 6.23 ± 0.23 μV, all p < 1e-4, Fig. 1e, Fig. S2). This aligns with the established understanding that N2/3 stages feature more pronounced slow-wave activity and robust spindles (Gaillard & Blois, 1981). Furthermore, we found that the amplitude of the SO DOWN-state (13.62 ± 0.67 μV) was significantly higher than the UP-state (12.98 ± 0.66 μV, t(106) = 4.82, p < 1e-4), consistent with previous work (Cash et al., 2009; Dang-Vu et al., 2008). Based on this, we used the DOWN-state to lock the timing of SOs for subsequent EEG-informed fMRI analysis. Descriptive information on sleep rhythms during these stages for all 107 subjects can be found in Table S2-S4.

SO-spindle coupling

Earlier version: SO-spindle coupling is considered important for sleep-dependent memory consolidation. In current study, we found that this coupling primarily occurs during N2/3 sleep stages (2.46 ± 0.06 times/min), significantly more than in N1 (2.23 ± 0.09 times/min, t(106) = 2.81, p = 0.0059) and REM stages (2.32 ± 0.09 times/min, t(106) = 1.99, p = 0.0492, Fig. 2b), consistent with previous findings (Ngo et al., 2013; Staresina et al., 2015). After extracting all EEG epochs where SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur during the transition to the UP state of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP state peak is 0°).

Revised version: SO-spindle coupling is considered important for sleep-dependent memory consolidation. In current study, we found that this coupling primarily occurs during N2/3 sleep stages (2.46 ± 0.06 times/min), significantly more than in N1 (2.23 ± 0.09 times/min, t(106) = 2.81, p = 0.0059) and REM stages (2.32 ± 0.09 times/min, t(106) = 1.99, p = 0.0492, Fig. 2b, Table S2-S4), consistent with previous findings (Ngo et al., 2013; Staresina et al., 2015). After extracting all EEG epochs where SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur slightly before the UP state peak of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP state peak is 0°).

Sleep rhythms and SO-spindle coupling.

Earlier version: a, The SO-spindle coupling in the temporal frequency domain. The upper two rows illustrates the spindle (12-16 Hz) phase-locked in the transition to UP-state of SO (0.16-1.25 Hz). The bottom row shows the averaged temporal frequency pattern across all instances of SO-spindle coupling and over all subjects. b, SO-spindle coupling density across different sleep stages. c, Differences between coupled and uncoupled sleep rhythms. The left panel shows the difference in amplitude between spindles coupled with SOs (Coupling) and spindles not coupled with SOs (Other). The right panel displays the difference in amplitude between SOs coupled with spindles (Coupling) and SOs not coupled with spindles (Other). d, Phase modulation of SO-spindle coupling. Spindle peaks cluster toward the UP-state of SO (i.e., 0°), where –π/2 reflects the transition from DOWN to UP-state. The histogram represents the distribution of coupling directions across all subjects, with the red line showing the mean. Coupling phases for each subject are plotted on the circle, with coupling strength color-coded. e, Distribution of spindle peaks on the SO phase during all SO-spindle coupling events across participants. The distribution is represented by a probability density function, and the density is evaluated at 100 equally spaced points covering the data range. Each dot represents data from an individual subject. Error bars indicate the SEM. * p < 0.05, ** p < 0.01, *** p < 0.001.

Revised version: a, The SO-spindle coupling in the temporal frequency domain. The upper two rows illustrates the spindle (12-16 Hz) phase-locked in the transition to UP-state of SO (0.16-1.25 Hz). The bottom row shows the averaged temporal frequency pattern across all instances of SO-spindle coupling and over all subjects. b, SO-spindle coupling density across different sleep stages. c, Differences between coupled and uncoupled sleep rhythms. The left panel shows the difference in amplitude between spindles coupled with SOs (Coupling) and spindles not coupled with SOs (Other). The right panel displays the difference in amplitude between SOs coupled with spindles (Coupling) and SOs not coupled with spindles (Other). d, Phase modulation of SO-spindle coupling. Spindle peaks cluster slightly before the UP-state peak of SO (i.e., 0°), where –π/2 reflects the transition from DOWN to UP-state. The histogram represents the distribution of coupling directions across all subjects, with the red line showing the mean. Coupling phases for each subject are plotted on the circle, with coupling strength color-coded. e, Distribution of spindle peaks on the SO phase during all SO-spindle coupling events across participants. The distribution is represented by a probability density function, and the density is evaluated at 100 equally spaced points covering the data range. Each dot represents data from an individual subject. Error bars indicate the SEM. * p < 0.05, ** p < 0.01, *** p < 0.001.

Interestingly, spindle events coupled with SOs had higher amplitudes (7.69 ± 0.26 μV) compared to those not coupled (7.33 ± 0.25 μV, t(106) = 4.65, p < 1e-4, Fig. 2c, left). Similarly, SO events coupled with spindles exhibited higher amplitudes (26.28 ± 1.34 μV) than those not coupled (24.56 ± 1.33 μV, t(106) = 6.86, p < 1e-4, Fig. 2c, right). These findings suggest coordinated neural processing during SO-spindle coupling. Next, we explore the brain-wide activation associated with these sleep rhythms.

Sleep rhythms associated brain activation

Earlier version: To investigate brain-wide activity during sleep rhythms, we modeled the fMRI blood oxygen level-dependent (BOLD) signal based on EEG-identified timing of SOs and spindle events during N2/3 sleep stages. The design matrix included the main effects of SOs, spindles, and their coupling (Fig. 3a, detailed in Methods). This method is validated in our previous work to effectively capture replay-aligned brain activation during rest (Huang et al., 2024). We found that SOs were associated with positive activation in the thalamus (ROI analysis, t(106) = 2.38, p = 0.0096, one-sample t test) and significant deactivation in the neocortex (Fig. 3b, see also Fig. S2), particularly in the default mode network (DMN), which includes the mPFC and posterior cingulate cortex (PCC), consistent with previous research (Dang-Vu et al., 2008). This deactivation of the neocortex, particularly the DMN, at the trough of SOs is intriguing, given that the DMN is thought to encode the internal model of the world or cognitive map (Baldassano, Hasson, & Norman, 2018; Constantinescu, O’Reilly, & Behrens, 2016; Park, Miller, & Boorman, 2021).

Revised version: To investigate brain-wide activity during sleep rhythms, we modeled the fMRI blood oxygen level-dependent (BOLD) signal based on EEG-identified timing of SOs and spindle events during N2/3 sleep stages. The design matrix included the main effects of SOs, spindles, and their coupling (Fig. 3a, detailed in Methods). This method is validated in our previous work to effectively capture replay-aligned brain activation during rest (Huang et al., 2024). We found that SOs were associated with positive activation in the thalamus (ROI analysis, t(106) = 2.38, p = 0.0096, one-sample t test) and significant deactivation in the neocortex (Fig. 3b, see also Fig. S5), particularly in the default mode network (DMN), which includes the mPFC and posterior cingulate cortex (PCC), consistent with previous research (Dang-Vu et al., 2008; Picchioni et al., 2011; Tagliazucchi et al., 2013). This deactivation of the neocortex, particularly the DMN, at the trough of SOs is intriguing, given that the DMN is thought to encode the internal model of the world or cognitive map (Baldassano, Hasson, & Norman, 2018; Constantinescu, O’Reilly, & Behrens, 2016; Park, Miller, & Boorman, 2021).

Brain-wide activation associated with sleep rhythms.

Earlier version: a, Simultaneous EEG-fMRI analysis framework for detecting brain-wide activation during sleep rhythms. Detected SOs, spindles, and their coupling were convolved with the hemodynamic response function (HRF) and downsampled to match fMRI temporal resolution. These events formed the design matrix for the general linear model (GLM) analysis of fMRI activity during sleep, linking the EEG-derived timing of sleep rhythms to the corresponding brain responses in fMRI. b, Brain-wide activation associated with SOs. The upper row illustrates SOs, and the lower row shows the fMRI activation pattern during SO events, whole-brain family-wise error (FWE) corrected at the cluster level (p < 0.05) with a cluster-forming voxel threshold of p < 0.001. c, Brain-wide activation associated with spindles. Same as panel b, but for spindle events. d, Brain-wide activation associated with SO-spindle coupling (compared to non-coupling events). e, Functional decoding using the ROI association method in Neurosynth. Each row corresponds to the brain-wide activation patterns for sleep rhythms shown in panels b-d, while each column corresponds to topics in the Neurosynth database (detailed in Methods). The left panel shows results using negative activation, and the right panel shows results using positive activation. Only topics with a decoded significance level of p < 0.05 are displayed.

Revised version: a, Simultaneous EEG-fMRI analysis framework for detecting brain-wide activation during sleep rhythms. Detected SOs, spindles, and their coupling were convolved with the hemodynamic response function (HRF) and downsampled to match fMRI temporal resolution. These events formed the design matrix for the general linear model (GLM) analysis of fMRI activity during sleep, linking the EEG-derived timing of sleep rhythms to the corresponding brain responses in fMRI. b, Brain-wide activation associated with SOs. The upper row illustrates SOs, and the lower row shows the fMRI activation pattern during SO events, whole-brain family-wise error (FWE) corrected at the cluster level (p < 0.05) with a cluster-forming voxel threshold of p < 0.001. c, Brain-wide activation associated with spindles. Same as panel b, but for spindle events. d, Brain-wide activation associated with SO-spindle coupling (compared to non-coupling events).

Earlier version: Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). For more detailed activation patterns, see Table S1-S3.

Revised version: Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). However, no decrease in DMN activity was found during SO-spindle coupling, and DMN activity during SO was significantly lower than during coupling (ROI analysis, t(106) = -4.17, p < 1e-4). For more detailed activation patterns, see Table S5-S7. We also varied the threshold used to detect SO events to assess its effect on hippocampal activation during SO-spindle coupling and observed that hippocampal activation remained significant when the percentile thresholds for SO detection ranged between 71% and 80% (see Fig. S6).

Earlier version: To understand the functional implications of these sleep rhythm-related brain activation, we performed open-ended cognitive state decoding using the NeuroSynth database (Yarkoni et al., 2011), examining associations between topic terms and regions of interest derived from the activation patterns of the sleep rhythms. Topic terms were ranked by their z-scores across the rhythms to assess overall functional relevance (see Methods for details).

Revised version: To understand the functional implications of these sleep rhythm-related brain activation, we performed open-ended cognitive state decoding using the NeuroSynth database (Yarkoni et al., 2011), examining associations between topic terms and regions of interest derived from the activation patterns of the sleep rhythms. Topic terms were ranked by their z-scores across the rhythms to assess overall functional relevance (see Methods for details). For positive activation patterns (Fig. S7a), SOs were associated with declarative (z = 6.31, p < 1e-4) and episodic memory (z = 5.08, p < 1e-4), while spindles aligned most with working memory (z = 4.08, p < 1e-4). SO-spindle coupling showed the strongest similarity to episodic (z = 3.66, p = 0.0003) and declarative memory (z = 2.21, p = 0.0271). For negative (deactivation) patterns (Fig. S7b), SOs were linked to declarative memory (z = 7.61, p < 1e-4), while spindles resembled task-related activations (z = 7.05, p < 1e-4). SO-spindle coupling had the highest similarity to task-related (z = 9.40, p < 1e-4) and working memory (z = 9.15, p < 1e-4). These findings suggest that internal memory-related regions are more likely to reactivate during SO-spindle coupling, whereas task-related activations, like working memory, are less prominent.

Removed in the revised version: For positive activation patterns (Fig. 3e, right), SOs were linked to patterns associated with declarative memory (z = 6.31, p < 1e-4) and episodic memory (z = 5.08, p < 1e-4), while spindles were most similar to activation patterns related to working memory (z = 4.08, p < 1e-4). SO-spindle coupling aligned most closely with activation patterns related to episodic memory (z = 3.66, p = 0.0003) and declarative memory (z = 2.21, p = 0.0271).

Removed in the revised version: For negative (deactivation) patterns (Fig. 3e, left), SOs were linked to patterns related to declarative memory (z = 7.61, p < 1e-4), while spindles were most similar to task-related activation patterns (z = 7.05, p < 1e-4). SO-spindle coupling showed the strongest similarity to patterns related to task-related activation (z = 9.40, p < 1e-4) and working memory (z = 9.15, p < 1e-4). These findings suggest that brain regions involved in internal memory processing are the most likely to reactivate during SO-spindle coupling events, while task-related activations, such as working memory, are the least likely to occur.

Functional connectivity during SO-spindle coupling

Previous studies have indicated that the hippocampal-cortical dialogue during sleep memory consolidation relies on the thalamus (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Latchoumane et al., 2017). To explore this in the human brain, we examined how functional connectivity (FC) changes when spindles couple (versus not) with SOs during sleep, using psychophysiological interaction (PPI) analysis (Friston et al., 1997). Building on the previous GLM (Fig. 3a), this model incorporated two additional regressors: the BOLD signal of the ROI seed and its interaction with SO-spindle events (Fig. 4a, detailed in Methods).

Functional connectivity changes during SO-spindle coupling.

a, The PPI analysis framework for detecting brain-wide connectivity changes during SO-spindle coupling. This starts by setting a specific ROI (e.g., the hippocampus) as the seed to extract the BOLD signal (physiological condition) and using identified SO-spindle coupling events as the psychological condition to compute the interaction term. The design matrix includes the main effects of the physiological and psychological conditions, along with their interaction. This analysis examines whether whole-brain communication with the hippocampus changes as a function of SO-spindle coupling. b, Hippocampus-based functional connectivity with the whole brain (main effect of hippocampus BOLD signal in PPI analysis). The hippocampus ROI is bilateral, anatomically defined (bottom, orange colour). Brain-wide connectivity is shown with whole-brain FWE correction at the cluster level (p < 1e-7) with a cluster-forming voxel threshold of punc. < 0.001 for visualization purpose. c, Same with panel b, but based on thalamus (bilateral anatomically defined ROI). d, Same with panel b, but based on the mPFC (bilateral functionally defined ROI, detailed in Methods). e, Functional connectivity changes during SO-spindle coupling for hippocampus-based (left bottom, orange colour), thalamus-based (middle, green colour), and mPFC-based (right bottom, purple colour) connectivity. The results of ROI analysis for each direction are shown on the arrows. * p < 0.05, ns., not significant. Abbreviations: FC - functional connectivity, PPI - psychophysiological interaction.

We first examined whole-brain FC (physiological term) in three ROIs: the hippocampus (bilateral, anatomically defined), thalamus (bilateral, anatomically defined), and mPFC (bilateral, functionally defined). The hippocampus-based FC was mainly connected with the DMN, particularly the mPFC (Fig. 4b). The thalamus-based FC was primarily linked to the basal ganglia and ACC (Fig. 4c), and the mPFC-based FC showed strong connections with the hippocampus and PCC (Fig. 4d).

Earlier version: Crucially, the interaction between FC and SO-spindle coupling revealed that only the hippocampus -> thalamus FC (ROI analysis, t(106) = 1.86, p = 0.0328) and thalamus -> mPFC FC (ROI analysis, t(106) = 1.98, p = 0.0251) significantly increased during SO-spindle coupling, with no significant changes in all other FC pathways (Fig. 4e). These findings suggest that the thalamus, likely through thalamic spindles, coordinate hippocampal-cortical communication during human NREM sleep.

Revised version: Crucially, the interaction between FC and SO-spindle coupling revealed that only the functional connectivity of hippocampus -> thalamus (ROI analysis, t(106) = 1.86, p = 0.0328) and thalamus -> mPFC (ROI analysis, t(106) = 1.98, p = 0.0251) significantly increased during SO-spindle coupling, with no significant changes in all other pathways (Fig. 4e). We also conducted PPI analyses for the other two events (SOs and spindles), and neither yielded significant connectivity changes in the three ROIs, as all failed to survive whole-brain FWE correction at the cluster level (p < 0.05). Together, these findings suggest that the thalamus, likely via spindles, coordinates hippocampal-cortical communication selectively during SO-spindle coupling, but not isolated SOs or spindle events alone.

Discussion

Earlier version: Using simultaneous EEG and fMRI during nocturnal sleep in 107 subjects, we identified robust coupling between sleep spindles and the (transition to) UP-state of SOs during N2 and N3 stages of NREM sleep. This coupling was linked to distinct neural activation patterns, including significant thalamic activation during spindle occurrences, reduced activity in the DMN during SOs, and elevated hippocampal activity during SO-spindle events. Moreover, we observed increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC during these SO-spindle events. Through cognitive state decoding, we found that brain-wide activation during SO-spindle events most closely resembles that of internal memory processing, while deactivation patterns align with those of working memory and task-related representation. These findings offer insights into the brain mechanisms and associated functional implications during human NREM sleep.

Revised version: Using simultaneous EEG and fMRI during nocturnal naps in 107 subjects, we identified robust coupling between spindles and the UP-state of SOs in deep NREM sleep (N2/3 stages). This coupling was linked to distinct neural activation patterns, including significant thalamic activation during spindle occurrences, reduced activity in the DMN during SOs trough, and selective hippocampal activation when spindles and SOs coupled. Furthermore, functional connectivity analyses revealed increased connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC, only during SO-spindle coupling, but not isolated SOs or spindle events alone. These results shed light on how the thalamus may coordinate hippocampal-cortical interactions during NREM sleep.

Added in the revised version: Our data provide insights into the neurobiological underpinnings of these sleep rhythms. SOs, originating mainly in neocortical areas such as the mPFC, alternate between DOWN- and UP-states. The thalamus generates sleep spindles, which in turn couple with SOs. Our finding that spindle peaks consistently occurred slightly before the UP-state peak of SOs (in 83 out of 107 participants), concurs with prior studies, including Schreiner et al. (2021). Yet it differs from some results suggesting spindles might peak right at the SO UP-state (Hahn et al., 2020; Helfrich et al., 2018). Such discrepancies could arise from differences in detection algorithms, participant age (Helfrich et al., 2018), or subtle variations in cortical-thalamic timing. Nonetheless, these results underscore the importance of coordinated SO-spindle interplay in supporting sleep-dependent processes.

Earlier version: We found that spindle amplitude peaks were nested within the transition to the UP-state of SOs in most subjects (83 out of 107), a coupling likely crucial for sleep-dependent memory consolidation. This is consistent with Schreiner et al. (2021), showing that memory reactivation is stronger when spindles occur closer to the SO UP-state. When modeling the timing of these sleep rhythms, we observed hippocampal activation during SO-spindle events, which is intriguing as it may indicate the presence of triple coupling, including hippocampal ripples. While our scalp EEG was not sensitive enough to detect hippocampal ripples—key markers of memory replay (Buzsáki, 2015)—recent intracranial EEG (iEEG) studies have observed ripples during human sleep, showing their coupling with both spindles (Ngo, Fell, & Staresina, 2020) and SOs (Staresina et al., 2015; Staresina et al., 2023). Therefore, the hippocampal involvement during SO-spindle events in our study may reflect memory replay from the hippocampus, propagated via thalamic spindles to distributed cortical regions.

Revised version: When modeling the timing of these sleep rhythms in the fMRI, we observed hippocampal activation selectively during SO-spindle events. This suggests the possibility of triple coupling (SOs–spindles– ripples), even though our scalp EEG was not sufficiently sensitive to detect hippocampal ripples—key markers of memory replay (Buzsáki, 2015). Recent iEEG evidence indicates that ripples often co-occur with both spindles (Ngo, Fell, & Staresina, 2020) and SOs (Staresina et al., 2015; Staresina et al., 2023). Therefore, the hippocampal involvement during SO-spindle events in our study may reflect memory replay from the hippocampus, propagated via thalamic spindles to distributed cortical regions.

Removed in the revised version: Our observation of reduced DMN activity during SOs is intriguing. SOs, generated in cortical networks, reach subcortical structures like the hippocampus and alternate between DOWN- and UP-states of synchronized neural activity (Nir et al., 2011). Due to the low temporal resolution of fMRI and the rhythmic pattern of SOs, it is difficult to distinguish the activation patterns of the DOWN- vs. UP-states (see also Fig. S2). It is likely that our findings of diminished DMN activity reflect brain activity during the SO DOWN-state, as this state consistently shows higher amplitude compared to the UP-state within subjects, which is why we modelled the SO trough as its onset in the fMRI analysis. The DMN is typically active during wakeful rest and is associated with self-referential processes like mind-wandering, daydreaming, and task representation (Yeshurun, Nguyen, & Hasson, 2021). Its reduced activity during SOs may signal a shift towards endogenous processes such as memory consolidation. Future research using techniques like iEEG, with both temporal and spatial precision, could clarify this possibility.

Earlier version: The thalamus, known to generate spindles (Halassa et al., 2011), plays a key role in producing and coordinating sleep rhythms (Coulon, Budde, & Pape, 2012; Crunelli et al., 2018), while the hippocampus is essential for memory consolidation during sleep, facilitating the transfer of memories to the neocortex for long-term storage (Buzsáki, 2015; Diba & Buzsáki, 2007; Singh, Norman, & Schapiro, 2022). The increased hippocampal and thalamic activity, along with strengthened functional connectivity between the hippocampus, thalamus, and mPFC during SO-spindle events, highlights the hippocampal-thalamic-neocortical information flow during human NREM sleep. This aligns with recent findings that suggest the thalamus orchestrates neocortical oscillations during sleep (Schreiner et al., 2022). These results emphasize the central role of thalamus in sleep-related memory consolidation, guiding the routing of information from the hippocampus to the neocortex (e.g., mPFC), a process critical to understanding memory consolidation during sleep.

Revised version: The thalamus, known to generate spindles (Halassa et al., 2011), plays a key role in producing and coordinating sleep rhythms (Coulon, Budde, & Pape, 2012; Crunelli et al., 2018), while the hippocampus is found essential for memory consolidation (Buzsáki, 2015; Diba & Buzsáki, 2007; Singh, Norman, & Schapiro, 2022). The increased hippocampal and thalamic activity, along with strengthened connectivity between these regions and the mPFC during SO-spindle events, underscores a hippocampal-thalamic-neocortical information flow. This aligns with recent findings suggesting the thalamus orchestrates neocortical oscillations during sleep (Schreiner et al., 2022). The thalamus and hippocampus thus appear central to memory consolidation during sleep, guiding information transfer to the neocortex, e.g., mPFC.

Added in the revised version: An intriguing aspect of our findings is the reduced DMN activity during SOs when modeled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP-from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.

Added in the revised version: To explore functional relevance, we employed an open-ended cognitive state decoding approach using meta-analytic data (NeuroSynth: Yarkoni et al. (2011)). Although this method usefully generates hypotheses about potential cognitive processes, particularly in the absence of a pre- and post-sleep memory task, it is inherently indirect. Many cognitive terms showed significant associations (16 of 50), such as “episodic memory,” “declarative memory,” and “working memory.” We focused on episodic/declarative memory given the known link with hippocampal reactivation (Diekelmann & Born, 2010; Staresina et al., 2015; Staresina et al., 2023). Nonetheless, these inferences regarding memory reactivation should be interpreted cautiously without direct behavioral measures. Future research incorporating explicit tasks before and after sleep would more rigorously validate these potential functional claims.

Added in the revised version: Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.

Added in the revised version: In summary, we found that SO-spindle coupling during deep NREM sleep is associated with distinct activation in both the thalamus and hippocampus, accompanied by increased connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC. These results emphasize the central role of the thalamus in orchestrating hippocampal-cortical interactions and highlight potential neural substrates underlying the complex interplay of SOs, spindles, and, possibly, ripples. Our open-ended decoding analyses suggest that these coordinated sleep events may be relevant for memory-related processes; however, future experiments incorporating explicit behavioral measures are necessary to confirm their functional significance. By revealing how these rhythms modulate widespread brain networks, our findings open avenues for developing targeted neuromodulatory strategies to enhance sleep-dependent cognitive functions and potentially address sleep-related disorders.

Removed in the revised version: Our open-ended cognitive state decoding analysis revealed that brain activation patterns linked to different sleep rhythms were associated with distinct cognitive functions. SOs and SO-spindle coupling were strongly connected to declarative and episodic memory, particularly during coupling events. In contrast, brain deactivation patterns were associated with task representations and working memory, reflecting the reduction in DMN activity (Baldassano, Hasson, & Norman, 2018; Smallwood et al., 2021). It is intriguing to speculate that during SOs, especially when aligned with spindle events, spontaneous replay of past episodic memories may occur, facilitating the consolidation of hippocampal memories into long-term declarative memory in the neocortex. This internal processing contrasts with the brain patterns associated with external tasks, such as working memory. This hypothesis warrants further studies to decode the content of sleep activity, capturing the replay of specific memories during sleep.

Removed in the revised version: Together, our findings provide insights into the neural mechanisms of sleep rhythms, emphasizing their coordination and the roles of the hippocampus and thalamus. They also highlight the importance of the role of thalamus in generating and coordinating sleep rhythms, the implications of reduced DMN activity during SOs, and the temporal dynamics of hippocampal-neocortical communication. Uncovering these brain-wide activation patterns is crucial for understanding the functional implications of sleep activity. These insights will lead to more targeted interventions that improve sleep quality and cognitive performance.

Methods

Participants & Protocol

A total of 138 healthy adults (age: 22.56 ± 0.30; 81 females) were recruited for the study. All participants had either normal vision or vision corrected to normal standards. None had any past psychiatric or neurological disorders. Prior to participation, they were screened for eligibility to undergo magnetic resonance imaging (MRI). The experiment received approval from Peking University ethics committee (reference number: #2015-26, #2018-11-05). Every participant provided written informed consent.

Earlier version: At the start of the experiment, subjects arrived early at the laboratory for registration. They signed informed consent forms, washed their hair, changed their clothes, and put on the EEG cap in preparation for the experiment. After these preliminary steps, participants entered the MRI scanning room by around 20:30, settled down for sleep, and data collection began. Given concern for potential data drift and decline in the signal-to-noise ratio due to extended and continuous fMRI scanning, the data acquisition was segmented into several sessions for overnight data collection. A “sleep” session was concluded either when a participant was fully awake (unable to fall asleep again or needed to use the restroom) or after 8192 seconds (the maximum scanning time of the scan, equating to 4096 volumes with 2 seconds TR).

Revised version: At the start of the experiment, subjects arrived early at the laboratory for registration. They signed informed consent forms, washed their hair, changed their clothes, and put on the EEG cap in preparation for the experiment. After these preliminary steps, participants entered the MRI scanning room by around 20:30, settled down for sleep, and data collection began. Given concern for potential data drift and decline in the signal-to-noise ratio due to extended and continuous fMRI scanning, the data acquisition was segmented into several sessions for data collection. A “sleep” session was concluded either when a participant was fully awake (unable to fall asleep again or needed to use the restroom) or after 8192 seconds (the maximum scanning time of the scan, equating to 4096 volumes with 2 seconds TR).

After removing subjects with excessive head movements (mean FD > 0.5) or those whose sleep durations were under 10 minutes, the final participant count for further analyses was 107 (age: 22.40 ± 0.34; 63 females) with a mean scanning duration of sleep of 3.01 ± 0.14 hours.

EEG data acquisition

Earlier version: EEG was recorded simultaneously with fMRI data using an MR-compatible EEG amplifier system (BrainAmps MR-Plus, Brain Products, Germany), along with a specialized electrode cap. The recording was done using 64 channels in the international 10/20 system, with the reference channel positioned at FCz. In order to adhere to polysomnography (PSG) recording standards, six electrodes were removed from the EEG cap: one for electrocardiogram (ECG) recording, two for electrooculogram (EOG) recording, and three for electromyogram (EMG) recording. EEG data was recorded at a sample rate of 5000 Hz, the resistance of the reference and ground channels was kept below 10 kΩ, and the resistance of the other channels was kept below 20 kΩ. To synchronize the EEG and fMRI recordings, the BrainVision recording software (BrainProducts, Germany) was utilized to capture triggers from the MRI scanner.

Revised version: EEG was recorded simultaneously with fMRI data using an MR-compatible EEG amplifier system (BrainAmps MR-Plus, Brain Products, Germany), along with a specialized electrode cap. The recording was done using 64 channels in the international 10/20 system, with the reference channel positioned at FCz. In order to adhere to polysomnography (PSG) recording standards, six electrodes were removed from the EEG cap: one for electrocardiogram (ECG) recording, two for electrooculogram (EOG) recording, and three for electromyogram (EMG) recording. EEG data was recorded at a sample rate of 5000 Hz, the resistance of the reference and ground channels was kept below 10 kΩ, and the resistance of the other channels was kept below 20 kΩ. To synchronize the EEG and fMRI recordings, the BrainVision recording software (BrainProducts, Germany) was utilized to capture triggers from the MRI scanner. The EEG sampling was synchronized to the MR scanner’s 10 MHz gradient system clock, ensuring a stable gradient artifact shape over time and enabling accurate artifact removal. This was achieved via the standard clock synchronization interface of the EEG amplifier, minimizing timing jitter and drift.

MRI data acquisition

Added in the revised version: All MRI data were acquired using a 20-channel head coil on a research-dedicated 3-Tesla Siemens Magnetom Prisma MRI scanner. Earplugs and cushions were provided for noise protection and head motion restriction. We chose the 20-channel head coil because it was compatible with our EEG system and ensured sufficient signal-to-noise ratio (SNR) for both EEG and fMRI acquisition. The EEG electrode cables were guided through the lateral and posterior openings of the head coil, secured with foam padding to reduce motion and minimize MR-related artifacts. Moreover, given the extended nature of nocturnal sleep recordings, the 20-channel coil helped maintain participant comfort while still achieving high-quality simultaneous EEG-fMRI data.

Earlier version: All MRI data were acquired using a 20-channel head coil on a research-dedicated 3-Tesla Siemens Magnetom Prisma MRI scanner. Earplugs and cushions were provided for noise protection and head motion restriction. For registration purposes, high-resolution T1-weighted anatomical images were acquired using a three-dimensional magnetization prepared rapid acquisition gradient-echo sequence (sequence specification: number of slices = 192, repetition time (TR) = 2,530 ms, echo time (TE) = 2.98 ms, inversion time = 1,100 ms, voxel size = 0.5 × 0.5 × 1 mm3, flip angle (FA) = 7°). The participants were asked to lie quietly in the scanner during data acquisition.

Revised version: For registration purposes, high-resolution T1-weighted anatomical images were acquired using a three-dimensional magnetization prepared rapid acquisition gradient-echo sequence (sequence specification: number of slices = 192, repetition time (TR) = 2,530 ms, echo time (TE) = 2.98 ms, inversion time = 1,100 ms, voxel size = 0.5 × 0.5 × 1 mm3, flip angle (FA) = 7°). The participants were asked to lie quietly in the scanner during data acquisition.

Earlier version: Then, the “sleep” session began after the participants were instructed to try and fall asleep. For the functional scans, whole-brain images were acquired using k-space and steady-state T2*-weighted gradient echo-planar imaging (EPI) sequence that is sensitive to the BOLD contrast. This measures local magnetic changes caused by changes in blood oxygenation that accompany neural activity (sequence specification: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm).

Revised version: Then, the “sleep” session began after the participants were instructed to try and fall asleep. For the functional scans, whole-brain images were acquired using k-space and steady-state T2*-weighted gradient echo-planar imaging (EPI) sequence that is sensitive to the BOLD contrast. This measures local magnetic changes caused by changes in blood oxygenation that accompany neural activity (sequence specification: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm). A relatively long TR and larger voxel size were chosen to optimize SNR and reduce susceptibility-related distortions, which are critical in EEG-fMRI sleep studies where head motion and physiological noise can be substantial. The longer TR allowed whole-brain coverage with sufficient temporal resolution, while the larger voxel size helped enhance BOLD sensitivity and minimize partial volume effects in deep brain structures (e.g., the thalamus and hippocampus), which are key regions of interest in this study.

EEG data preprocessing and sleep stage scoring

EEG data collected inside MRI scanner was contaminated by imaging, ballistocardiographic (BCG) and ocular artifacts. To deal with these artifacts, the preprocessing processes was divided into two parts: automated processing and manual processing.

Earlier version: The automated processing was carried out using custom Python scripts and the MNE toolkit (Gramfort et al., 2013). This encompassed the following procedures: I. The raw EEG data were downsampled from 5000 Hz to 2048 Hz. This downsampled rate accelerates subsequent analyses while retaining the necessary precision to distinguish MR noise. II. Imaging artifacts were removed using the Average Artefact Subtraction method (Allen et al., 2000). III. We then applied 0.1 - 30 Hz band-pass finite impulse response (FIR) filters to all channels. The EEG and EOG electrodes were re-referenced by subtracting the value from the M1 or M2 electrode (whichever was from the opposite hemisphere). The outer EMG electrodes were re-referenced to the middle one. IV. The signal space projection algorithm was applied to extract all epochs based on TRs and remove residual MR noise, as well as automatically extracting all heartbeat epochs using ECG electrode data for the removal of BCG noise. V. The data were further downsampled to 100 Hz and subjected to Independent Component Analysis (ICA).

Revised version: The automated processing was carried out using custom Python scripts and the MNE toolkit (Gramfort et al., 2013). This encompassed the following procedures: I. The raw EEG data were downsampled from 5000 Hz to 2048 Hz. This downsampled rate accelerates subsequent analyses while retaining the necessary precision to distinguish MR noise. II. Imaging artifacts were removed using the Average Artefact Subtraction method (Allen, Josephs, & Turner, 2000). III. We then applied 0.1 - 30 Hz band-pass finite impulse response (FIR) filters to all channels. The EEG and EOG electrodes were re-referenced by subtracting the value from the M1 or M2 electrode (whichever was from the opposite hemisphere). The outer EMG electrodes were re-referenced to the middle one. IV. The signal space projection algorithm was applied to extract all epochs based on TRs and remove residual MR noise, as well as automatically extracting all heartbeat epochs using ECG electrode data for the removal of BCG noise. V. The data were further downsampled to 100 Hz and subjected to Independent Component Analysis (ICA).

Following this, the manual processing is carried out with labelling and removing the physiological artifacts ICs belonging to eye, muscle and residual BCG through the use of custom MATLAB scripts and the EEGLAB toolkit (Delorme & Makeig, 2004). Sleep stages were then scored by YASA toolkit (Vallat & Walker, 2021) after preprocessing, necessitating data from electrodes C3, left EOG, and left EMG. The classification of N2/3 stages were further verified by experts for validity.

MRI data preprocessing

Results included in this manuscript come from preprocessing performed using fMRIPrep 21.0.4 (Esteban et al., 2020; Esteban et al., 2019), which is based on Nipype 1.6.1 (Gorgolewski et al., 2011). Many internal operations of fMRIPrep use Nilearn 0.8.1 (Abraham et al., 2014), mostly within the functional processing workflow. For more details of the pipeline, see https://fmriprep.readthedocs.io/en/latest/workflows.html.

Conversion of data to the brain imaging data structure standard

To facilitate further analysis and sharing of data, all study data were arranged according to the Brain Imaging Data Structure (BIDS) specification using dcm2bids tool, which is freely available from https://unfmontreal.github.io/Dcm2Bids/.

Anatomical data preprocessing

The T1-weighted (T1w) image was corrected for intensity non-uniformity (INU) with N4BiasFieldCorrection (Tustison et al., 2010), distributed with ANTs 2.3.3 (Avants et al., 2008), and used as T1w-reference throughout the workflow. The T1w-reference was then skull-stripped with a Nipype implementation of the antsBrainExtraction.sh workflow (from ANTs), using OASIS30ANTs as target template. Brain tissue segmentation of cerebrospinal fluid (CSF), white-matter (WM) and gray-matter (GM) was performed on the brain-extracted T1w using fast (FSL 6.0.5.1:57b01774, RRID:SCR_002823) (Zhang, Brady, & Smith, 2001). Volume-based spatial normalization to one standard space (MNI152NLin2009cAsym) was performed through nonlinear registration with antsRegistration (ANTs 2.3.3), using brain-extracted versions of both T1w reference and the T1w template. The following template was selected for spatial normalization: ICBM 152 Nonlinear Asymmetrical template version 2009c (Fonov, 2009) [RRID:SCR_008796; TemplateFlow ID: MNI152NLin2009cAsym].

Functional data preprocessing

First, a reference volume and its skull-stripped version were generated using a custom methodology of fMRIPrep. Head-motion parameters with respect to the BOLD reference (transformation matrices, and six corresponding rotation and translation parameters) are estimated before any spatiotemporal filtering using mcflirt (Jenkinson et al., 2002) [FSL 6.0.5.1:57b01774]. BOLD runs were slice-time corrected to 0.96s (0.5 of slice acquisition range 0s-1.92s) using 3dTshift from AFNI (Cox & Hyde, 1997). The BOLD time-series (including slice-timing correction when applied) were resampled onto their original, native space by applying the transforms to correct for head-motion. These resampled BOLD time-series will be referred to as preprocessed BOLD in original space, or just preprocessed BOLD. The BOLD reference was then co-registered to the T1w reference using mri_coreg (FreeSurfer) followed by flirt (Jenkinson & Smith, 2001) [FSL 6.0.5.1:57b01774] with the boundary-based registration (Greve & Fischl, 2009) cost-function. Co-registration was configured with six degrees of freedom. Several confounding time-series were calculated based on the preprocessed BOLD: framewise displacement (FD), DVARS and three region-wise global signals. FD was computed using two formulations following Power (absolute sum of relative motions) (Power et al., 2014) and Jenkinson (relative root mean square displacement between affines, Jenkinson et al. (2002)). FD and DVARS are calculated for each functional run, both using their implementations in Nipype (following the definitions by Power et al. (2014)). The three global signals are extracted within the CSF, the WM, and the whole-brain masks. Additionally, a set of physiological regressors were extracted to allow for component-based noise correction (CompCor) (Behzadi et al., 2007). Principal components are estimated after high-pass filtering the preprocessed BOLD time-series (using a discrete cosine filter with 128s cut-off) for the two CompCor variants: temporal (tCompCor) and anatomical (aCompCor). tCompCor components are then calculated from the top 2% variable voxels within the brain mask. For aCompCor, three probabilistic masks (CSF, WM and combined CSF+WM) are generated in anatomical space. The implementation differs from that of Behzadi et al. (2007) in that instead of eroding the masks by 2 pixels on BOLD space, the aCompCor masks are subtracted a mask of pixels that likely contain a volume fraction of GM. This mask is obtained by thresholding the corresponding partial volume map at 0.05, and it ensures components are not extracted from voxels containing a minimal fraction of GM. Finally, these masks are resampled into BOLD space and binarized by thresholding at 0.99 (as in the original implementation). Components are also calculated separately within the WM and CSF masks. For each CompCor decomposition, the k components with the largest singular values are retained, such that the retained components’ time series are sufficient to explain 50 percent of variance across the nuisance mask (CSF, WM, combined, or temporal). The remaining components are dropped from consideration. Frames that exceeded a threshold of 0.5 mm FD or 1.5 standardised DVARS were annotated as motion outliers.

The BOLD time-series were resampled into standard space, generating a preprocessed BOLD run in MNI152NLin2009cAsym space. First, a reference volume and its skull-stripped version were generated using a custom methodology of fMRIPrep. All resamplings can be performed with a single interpolation step by composing all the pertinent transformations (i.e., head-motion transform matrices, susceptibility distortion correction when available, and co-registrations to anatomical and output spaces). Gridded (volumetric) resamplings were performed using antsApplyTransforms (ANTs), configured with Lanczos interpolation to minimize the smoothing effects of other kernels (Lanczos, 1964). Non-gridded (surface) resamplings were performed using mri_vol2surf (FreeSurfer).

EEG sleep rhythms analysis

Earlier version: In line with other studies examining the coordination between SOs and spindles, especially the spatial distribution of these sleep rhythms (Massimini et al., 2004; Molle et al., 2011), EEG data from the F3 electrode was used to quantify the occurrence of SOs and spindles. Using the established detection algorithms (detailed below), we identified SOs and sleep spindles for each subject.

Revised version: In line with other studies examining the coordination between SOs and spindles, particularly their spatial distribution (Massimini et al., 2004; Molle et al., 2011), EEG data from the F3 electrode were used to quantify the occurrence of SOs and spindles. Using the established detection algorithms (detailed below), we identified SOs and sleep spindles for each subject.

Detection of SOs

Earlier version: Data underwent a 0.16-1.25 Hz bandpass filtering (Butterworth filter, order 3, bidirectional filtering for zero-phase). After identifying all positive-to-negative zero crossings, potential SOs were determined based on the duration between consecutive zero crossings (ranging from 0.8 s to 3 s). We computed the amplitude range (peak minus trough) for each potential SO. Only those surpassing the 75th percentile in amplitude were marked as SOs (Schreiner et al., 2021). The onsets and durations of these SOs were saved.

Revised version: Data were first bandpass-filtered between 0.16 and 1.25 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). After identifying all positive-to-negative zero crossings, potential SOs were defined based on the interval between consecutive zero crossings (ranging from 0.8 s to 3 s). For each potential SO, we calculated the amplitude range (peak minus trough). Only those exceeding the 75th percentile in amplitude were labeled as SOs, following previous work (Schreiner et al., 2021).

Detection of sleep spindles

Earlier version: Data underwent a 12-16 Hz bandpass filtering (Butterworth filter, order 3, bidirectional filtering for zero-phase). The root mean square (RMS) of the filtered signal was calculated with a sliding time window of 200 ms. The amplitude threshold for spindles was set at the 75th percentile of RMS values. RMS values exceeding this threshold and lasting between 0.5 s and 3 s were identified as spindles (Staresina et al., 2015). We saved all their onsets and durations.

Revised version: Data were bandpass-filtered between 12 and 16 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). The root mean square (RMS) of the filtered signal was computed with a 200 ms sliding time window. The amplitude threshold for spindles was set at the 75th percentile of RMS values. Any RMS values that surpassed this threshold for 0.5 s to 3 s were identified as spindles (Staresina et al., 2015).

Detection of SO-spindle couplings

Earlier version: From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO range, we checked for a spindle peak. If detected, it was recorded as an SO-spindle coupling event. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough as the central reference, following Schreiner et al. (2021). Data were extracted in a window of [-4 s to 4 s] around this point, forming the epoch for each coupling event. All SO-spindle epochs and their associated spindle peak times were saved.

Revised version: From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO interval, we checked for a spindle peak; if found, it was recorded as an SO-spindle coupling event. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough (as the central reference), following Schreiner et al. (2021). We extracted data in a [-4 s to 4 s] window around this point, forming the epoch for each coupling event.

Added in the revised version: We note that the above methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM (see also Figure S2-S4, Table S1-S4).

Added in the revised version: Event-related potentials (ERP) analysis

Added in the revised version: After completing the detection of each sleep rhythm event, we performed ERP analyses for SOs, spindles, and coupling events in different sleep stages. Specifically, for SO events, we took the trough of the DOWN-state of each SO as the zero-time point, then extracted data in a [-2 s to 2 s] window from the broadband (0.1–30 Hz) EEG and used [-2 s to -0.5 s] for baseline correction; the results were then averaged across 107 subjects (see Fig. S2a). For spindle events, we used the peak of each spindle as the zero-time point and applied the same data extraction window and baseline correction before averaging across 107 subjects (see Fig. S2b). Finally, for SO-spindle coupling events, we followed the same procedure used for SO events (see Fig. 2a, Figs. S3–S4).

Time-frequency analysis

We performed this analysis on the SO-spindle coupling epochs using the FieldTrip toolbox (Oostenveld et al., 2011). Fourier analysis was executed on the data, utilizing a sliding time window with a forward shift of 50 ms. The window length was defined as five cycles of the present frequency (frequency range: 1-30 Hz, step size: 1 Hz). Before Fourier analysis, windowed data segments underwent multiplication by a Hanning taper. Subsequently, power values were normalized to z-scores within the time range of [-4 s to 4 s], selecting a longer duration to prevent edge artifacts from low-frequency activity in the target time window of [-1.5 s to 1.5 s].

Preferred phase analysis

All SO-spindle coupling epochs were filtered to the SO frequency band (0.16-1.25 Hz, 3rd-order Butterworth filter), and the instantaneous SO phase was determined through the Hilbert transform. Using the saved spindle peak times for each SO-spindle coupling, the preferred phase was deduced. The distribution of these phases across participants was then tested for uniformity using the Rayleigh test.

fMRI GLM analysis

All fMRI GLM analyses were conducted using only data from the N2/3 stages of NREM sleep. Segments of sleep data from these stages were concatenated for each participant. After pre-processing, the fMRI data was smoothed with an 8 mm FWHM kernel. All images underwent a temporal high-pass filter (width of 100 s), and the autocorrelation of the hemodynamic response was modelled via an AR(1) model. We then conducted General Linear Model (GLM) analyses to identify whole-brain activations associated with sleep rhythms and their coupling. We incorporated nuisance regressors estimated during the preprocessing with fMRIprep: six rigid-body motion-correction parameters identified during realignment (comprising three translation and rotation parameters), the mean White Matter, and the mean Cerebral Spinal Fluid signals. The cosine drift model was employed to detect and eliminate data fluctuations, thereby enhancing the accuracy of actual experimental effects. The GLMs were designed to evaluate the influence of experimental conditions on BOLD responses. All whole-brain analyses were thresholded using whole-brain FWE corrected at the cluster level (p < 0.05), with a cluster-forming voxel threshold of punc. < 0.001, unless otherwise specified.

The fMRI GLM analysis was designed to look for brain activations pertinent to SO, spindle, and SO-spindle coupling. Drawing from the EEG results concerning the timing of SO and spindle, we identified the trough of the SO as the onset for SO events and the peak of the spindle as the onset for spindle events. We chose the trough of the SO and the peak of the spindle as the onset points, primarily due to the larger amplitudes observed at these moments (thereby more robust to noise). This was also the common choice in previous EEG studies on SO and spindles (Ngo, Fell, & Staresina, 2020; Staresina et al., 2015; Staresina et al., 2023). To model SO-spindle coupling in fMRI, we followed Schreiner et al. (2021), and defining SO-spindle complex as the peak of spindle appeared within 1.5 seconds after SO trough. When such a peak was found, that particular SO was labelled as an SO-spindle event, with its SO trough labelled as the event onset. After convolving these events with the canonical Hemodynamic Response Function (HRF) of SPM and downsampling them to align with the fMRI time resolution, they were integrated as regressors in the design matrix for GLM analysis, allowing us to determine the effects of SO and spindle, as well as the SO-spindle coupling (Fig. 3b-d).

Earlier version: We have also conducted a control analysis aimed to distinguish brain activations associated with DOWN- vs. UP state of SOs. Drawing again from EEG data on SO and spindle, we randomly selected half of the SO events using the peak as the onset and the other half using the trough, to reduce collinearity between the two regressors. As with the main analysis, we identified the spindle peak as the onset for spindle events (Fig. S2).

Revised version: We have also conducted a control analysis aimed to distinguish brain activations associated with DOWN- vs. UP state of SOs. Drawing again from EEG data on SO and spindle, we randomly selected half of the SO events using the peak as the onset and the other half using the trough, to reduce collinearity between the two regressors. As with the main analysis, we identified the spindle peak as the onset for spindle events (Fig. S5).

ROI analysis

The objective of our ROI analysis was to test the activation of specific brain regions during distinct sleep rhythms. Within each ROI, the beta values at the subject-level were averaged across all voxels to facilitate subsequent statistical inferences (one-sample t test). We defined the mDMN and mPFC ROIs functionally based on the meta-analysis results in Neurosynth (Yarkoni et al., 2011), by the keyword “medial default mode network” and “ventromedial prefrontal cortex” respectively. In addition to mDMN and mPFC, we defined other ROIs anatomically, using the high-resolution Harvard-Oxford Atlas probabilistic atlas (Desikan et al., 2006).

Open-ended Cognitive State Functional Decoding

For functional decoding, we followed the approach in Margulies et al. (2016). We used the NeuroSynth (Yarkoni et al., 2011) meta-analytic database (www.neurosynth.org) to assess topic terms associated with the sleep rhythms activation patterns. ROI masks were created from the whole-brain activation patterns associated with the three sleep rhythms in Fig. 3b-d. These activation patterns were binarized using a threshold of p < 0.001, forming a brain mask for each sleep rhythm.

For each ROI mask, the analysis output was a z-statistic associated with a feature term, and the terms were ranked based on the weighted sum of their z-scores. The feature terms were derived from the most general set of 50 topic terms available in the NeuroSynth v3 dataset (https://github.com/NeuroanatomyAndConnectivity/gradient_analysis/blob/master/gradient_data/neurosynth/v3-topics-50-keys.txt).

Earlier version: From the 50 topics, 16 exceeded the threshold of z > 1.96 (p < 0.05). Two terms were removed as they did not represent coherent cognitive functions, leaving 14 relevant topic terms (Fig. 3e). These included: episodic memory, declarative memory, working memory, task representation, language, learning, faces, visuospatial processing, category recognition, cognitive control, reading, cued attention, inhibition, and action.

Revised version: From the 50 topics, 16 exceeded the threshold of z > 1.96 (p < 0.05). Two terms were removed as they did not represent coherent cognitive functions, leaving 14 relevant topic terms (Fig. S7). These included: episodic memory, declarative memory, working memory, task representation, language, learning, faces, visuospatial processing, category recognition, cognitive control, reading, cued attention, inhibition, and action.

PPI analysis

Earlier version: PPI (psychophysiological interaction) analysis is designed to quantify context-dependent alterations in effective connectivity between brain regions. For instance, it can determine whether hippocampus-based functional connectivity with the thalamus varies based on the occurrence of an SO-spindle coupling event. We conducted three whole-brain PPI analyses with nilearn to investigate interactions based in the hippocampus, thalamus or mPFC during SO-spindle coupling, utilizing the same ROIs.

Revised version: PPI (psychophysiological interaction) analysis is designed to quantify context-dependent alterations in effective connectivity between brain regions. For instance, it can determine whether hippocampus-based functional connectivity with the thalamus varies based on the occurrence of an SO-spindle coupling event. We conducted nine whole-brain PPI analyses with nilearn to investigate interactions based in the hippocampus, thalamus or mPFC during SO-spindle coupling, isolated SO and isolated spindle, utilizing the same ROIs. Both the main effects of physiological and physiological condition, as well as their interaction were modelled in the design matrix for PPI analysis. The main effect of physiological term produced functional connectivity maps of the entire brain in relation to the specified ROI (Fig. 4b-d). While the interaction term asked for the change of functional connectivity along with the SO-spindle coupling (Fig. 4e).

Removed in the revised version: Both the main effects of physiological and physiological condition, as well as their interaction were modelled in the design matrix for PPI analysis. The main effect of physiological term produced functional connectivity maps of the entire brain in relation to the specified ROI (Fig. 4b, c, d). While the interaction term asked for the change of functional connectivity along with the SO-spindle coupling (Fig. 4e).